# SLU8 Object Oriented Programming | Inheritance: Exercise Notebook

***

### Part 1
#### Exercise 1

Create a class named `Vehicle()` and define the `__init__()` method of the created class to take four inputs:

1. self
2. model
3. color
4. year
5. engine_type


Now, do the following:
- Assign these inputs to member variables of the same name by using the self keyword.

- Define an extra variable in the `__init__()` method, without user input, called `condition` and with default value of **"new"**.

- Create a `drive_it()` method that will override the vehicle condition to the string **"used"**

In the end, create the `my_car` object for a **1985** **silver** **"DeLorean"** with a **combustion** engine



<!-- grade_id: cell-152d925316906cc7, points: 0, grade: false, solution: true -->
```python
class Vehicle:
   
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-d6fb846dc8915dc9, points: 2, grade: true, solution: false -->
```python
assert my_car.model == 'DeLorean'
assert my_car.color == 'silver'
assert my_car.year == 1985
assert my_car.engine_type == 'combustion'
assert my_car.condition == 'new'

print("""If my calculations are correct, you are a proud owner of a brand %s, %s %s %s with a %s engine!
That's heavy!""" % (my_car.condition,
                    my_car.year,
                    my_car.color,
                    my_car.model,
                    my_car.engine_type))
```

<img src="./assets/delorean.jpg" width="400"/>

***

#### Exercise 2
Now just go and take you brand new car for a spin, using the `drive_it()` method of your `my_car` object.

<!-- grade_id: cell-49ead50591e8361d, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

# after driving it for the first time, you vehicle condition will be updated:
print(my_car.condition)
```

<!-- grade_id: cell-c48e251cfd4b31e2, points: 1, grade: true, solution: false -->
```python
assert my_car.condition == 'used', "The condition of your car is not correct. Make sure you use the 'drive_it()' method"
print("I am sorry to inform you that the condition of your beloved DeLorean is now " + my_car.condition + "!")
```

#### Exercise 3

As seen before, one of the benefits of classes is that we can create more complicated classes that inherit variables or methods from their **parent classes**. This saves us time and helps us build more complicated objects, since these **child classes** can also include additional variables or methods.
+ Create an `EletricalVehicle` class that **inherits** from `Vehicle`;
+ Create a new method `battery_level()` that returns a (pseudo)random integer from 1 to 100 *
+ Then, create an eletrical vehicle named `my_tesla`, for a **2020 black Model_S** with an engine type of **molten salt battery**

Make sure that `my_tesla` belongs to `EletricalVehicle` Class.

***** the `random` package is already imported in the first line, it has function `randint()` for generating pseudo-random integers.

<!-- grade_id: cell-9d9a01b72fa64213, points: 0, grade: false, solution: true -->
```python
import random

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7722214f55e01cee, points: 2, grade: true, solution: false -->
```python
# __init__() tests:
assert my_tesla.model == 'Model_S'
assert my_tesla.color == 'black'
assert my_tesla.year == 2020
assert my_tesla.engine_type == 'molten salt battery'
assert my_tesla.condition == 'new'

# battery level tests:
assert (my_tesla.battery_level() >= 1) & (my_tesla.battery_level() <= 100)

# isistance tests:
assert (isinstance(my_tesla, EletricalVehicle)) == True, 'my_tesla must belong to the EletricalVehicle Child Class'
assert (isinstance(my_tesla, Vehicle)) == True, 'my_tesla must belong to the Vehicle Parent Class'
assert (isinstance(my_car, Vehicle)) == True
assert (isinstance(my_car, EletricalVehicle)) == False

print("""Woah, Nikolas would be proud!
A %s, %s %s %s with a %s!
Again...That's heavy!""" % (my_tesla.condition,
                    my_tesla.year,
                    my_tesla.color,
                    my_tesla.model,
                    my_tesla.engine_type))
print('')
print(isinstance(my_tesla, EletricalVehicle),'my_tesla belongs to the EletricalVehicle Child Class.')
print(isinstance(my_tesla, EletricalVehicle),'my_tesla belongs to the Vehicle Parent Class.')
```

<img src="./assets/tesla_edison_jk.jpg" width="400"/>

***

### Part 2
#### Exercise 4
Start by creating a new class `Rectangle` with the following:
+ Attributes:
    - **length**
    - **width**
+ Methods:
    - `area()`
    - `perimeter()`

<!-- grade_id: cell-e1f05519ef39b0f6, points: 0, grade: false, solution: true -->
```python
# class Rectangle():
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c0da80882a9d38bc, points: 2, grade: true, solution: false -->
```python
# considering a 10,50 m by 20,50 m rectangle:
r = Rectangle(10.50, 20.50)
assert r.area() == 215.25, 'Area computation is not correct!'
assert r.perimeter() == 62.0, 'Perimeter computation is not correct'
```

#### Exercise 5
Create a `Square` class, as a child of the `Rectangle` class. Change only the `__init__()` method in a way that now it only accepts one argument 'side'. You can call the init method from the parent class and assign both length and width to be equal to 'side'

(hint: use the `super()` method)

<!-- grade_id: cell-9262f02e916a66f1, points: 0, grade: false, solution: true -->
```python
# class Square
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0d959343e7bd5637, points: 1, grade: true, solution: false -->
```python
# considering a 10,50 m by 10,50 m square:
s = Square(10.50)

assert s.area() == 110.25, 'Area computation is not correct!'
assert s.perimeter() == 42.0, 'Perimeter computation is not correct'
```

#### Exercise 6

Create a new class `Triangle`:
+ Attributes:
    - **base**
    - **height**
+ Methods:
    - `tri_area()` which calculates the area of the triangle (do you remember the formula? :D)

<!-- grade_id: cell-361d990792a542f3, points: 0, grade: false, solution: true -->
```python
# class Triangle
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a0f3c54e307aeee9, points: 1, grade: true, solution: false -->
```python
t = Triangle(10, 20)

assert t.tri_area() == 100, 'The area of the triangle is miscalculated!'

print("The area, in cm2, of our triangle is ", t.tri_area())
```

#### Exercise 7

Now for the `RightPyramid` class, create methods that compute the area of this geometric solid using both the `Square` and `Triangle` classes as parents.

Override the `area()` method in order to compute the equation of the RightPyramid area. **Hint:** use the `super()` function to  make everything easier.

<img src="assets/pyramid.png" width="400"/>

+ Attributes:
    - **base**
    - **slant_height**
+ Methods:
    - `area()`

<!-- grade_id: cell-18e78aa00977daf8, points: 0, grade: false, solution: true -->
```python
# class RightPyramid
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7a3b92fb1669680d, points: 3, grade: true, solution: false -->
```python
rp = RightPyramid(10, 20)

print(rp.area())

assert isinstance(rp, RightPyramid)
assert isinstance(rp, Square), 'Probably your class is not properly defined, should be a child class of the Square classe'
assert isinstance(rp, Triangle), 'Probably your class is not properly defined, should be a child class of the Triangle classe'
assert rp.area() == 500, 'The area of the right pyramid is miscalculated!'

```

***

### Part 3

In the following exercice the goal it is to create a BankAccount program, composed by one parent class and two child classes, in order to manage the clients accounts, either through the main (current) account or the savings account.

This is the general idea:

<img src="assets/bank_account.jpg" width="400"/>



#### Exercise 8

First things first, the `BankAccount` class:

+ Attributes:
    - **costumer**
    - **account_number**
    - **balance** with default value of 0 (zero)
+ Methods:
    - `deposit()`: deposit **amount** into the bank account. (increase the balance). If the **amount** is invalid (< 0) it should print an alert message to the user.
    - `withdraw()`: withdraw **amount** from the bank account. (decrease the balance). This method must ensure that the account has sufficient funds for the asked **amount**. Again, it should be > 0. Print an alert if the value is less than zero or if it's bigger than the balance
    - `check_balance()`: print a statement of the account balance, example _"The balance of account number 48522888 is €1000"_

<!-- grade_id: cell-aec63d5088ce5f30, points: 0, grade: false, solution: true -->
```python
class BankAccount:
    """
    Abstract base class representing a bank account
    """
    currency = '€'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-5fb3aa9dc9bb6e4c, points: 2, grade: true, solution: false -->
```python
# Start by creating a new account:
my_account = BankAccount('John Doe', 21457288)


# assert the class:
assert isinstance(my_account, BankAccount) == True, 'Your class is not properly defined'

# assert the properties of the class:
assert my_account.account_number == 21457288, 'You have to define an "account_number" attribute for your class!'
assert my_account.costumer == 'John Doe', 'You have to define a "costumer" attribute for your class!'
assert my_account.balance == 0, 'The default value of "balance" should be 0!'
assert my_account.currency == '€', 'Currency symbol is not working... Did you change the existing line?'

print("Nice, the account for %s was successfully created. His account number is %s and for now is account is empty. \n"
     % (my_account.costumer, my_account.account_number))

# Lets deposit that big, fat check !
my_account.deposit(25000)

assert my_account.balance == 25000, 'Your balance, after making a deposit should change!'

print("%s just made a deposit of €25'000. His new account balance is: %s%s \n"
      %(my_account.costumer, my_account.currency, my_account.balance) )

# And then just go shopping (online... better be online nowadays):
my_account.withdraw(2250)

assert my_account.balance == 22750, 'Your balance, after making a withdrawal should change or it is not correct!'
```

<!-- grade_id: cell-aaeb81da6d6cc191, points: 1, grade: true, solution: false -->
```python
# Finally let's ensure that your program does not accept invalid operations:
my_account.deposit(-10)
my_account.withdraw(-10)
assert my_account.balance == 22750, 'Your balance, after making an invalid deposit should not change!'
```

#### Exercise 9

Next we need to develop the `SavingsAccount` class, as a child class for `BankAccount`:

+ Attributes:
    - *interest_rate*
    - All the BankAccount attributes
    
+ Methods:
    - `add_interest()`: add interest to the account at the rate *interest_rate*, thus updating the account balance.

<!-- grade_id: cell-8f1a5436b1aa5fd0, points: 0, grade: false, solution: true -->
```python
# class SavingsAccount...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2e5fd33c1f193d5a, points: 2, grade: true, solution: false -->
```python
# Creating a new savings account:
my_savings = SavingsAccount('Jane Doe', 41522887, 0.55, 10000)

# assert the class:
assert isinstance(my_savings, SavingsAccount) == True, 'Your class is not properly defined'
assert isinstance(my_savings, BankAccount) == True, 'Your class has no relation with the BankAccount class'

# assert the properties of the class:
assert my_savings.account_number == 41522887, 'You have to define an "account_number" attribute for your class!'
assert my_savings.costumer == 'Jane Doe', 'You have to define a "costumer" attribute for your class!'
assert my_savings.balance == 10000, 'The default value of "balance" should be 0!'
assert my_savings.currency == '€', 'Currency symbol is not working... Did you changed the existing line?'

print("Nice, the account for %s was successfully created. Her account number is %s. \n"
     % (my_savings.costumer, my_savings.account_number))

# Check account balance:
print("Checking account balance before interest:")
my_savings.check_balance()

# adding some interest to the existing account:
my_savings.add_interest()
print("\nChecking account balance after interest:")
my_savings.check_balance()

assert my_savings.balance == 10055, "Maybe the created savings account is not generating any interest!"
```

#### Exercise 10

Finally we need to implement the `CurrentAccount` class, as a child class for `BankAccount`, that has a similar structure of the previous exercise:

+ Attributes:
    - **annual_fee**
    - **transaction_limit**
    - All the BankAccount attributes
    
+ Methods:
    - `withdraw()`: override the parent class method and create an if clause with an extra condition, in order to limit the possible amount that a person can withdraw from a Current Account daily. **Hint:** use the *transaction_limit* variable you just initialized.
    - `apply_annual_fee()`: Deduct the annual fee from the account balance, making sure that the account never goes negative! **Hint:** use python's `max()` function

<!-- grade_id: cell-6d21acc19869fec0, points: 0, grade: false, solution: true -->
```python
# class CurrentAccount...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4f6963493a0e8fe9, points: 1, grade: true, solution: false -->
```python
# create a new current account:
my_current = CurrentAccount('Richard Roe', 78300991, 20., 200.)

# assert the class:
assert isinstance(my_current, CurrentAccount) == True, 'Your class is not properly defined'
assert isinstance(my_current, BankAccount) == True, 'Your class has no relation with the BankAccount class'

# assert the properties of the class:
assert my_current.account_number == 78300991, 'You have to define an "account_number" attribute for your class!'
assert my_current.costumer == 'Richard Roe', 'You have to define a "costumer" attribute for your class!'
assert my_current.balance == 0, 'The default value of "balance" should be 0!'
assert my_current.currency == '€', 'Currency symbol is not working... Did you changed the existing line?'

print("Nice, the account for %s was successfully created. Her account number is %s and for now is account is empty.\n"
     % (my_savings.costumer, my_savings.account_number))
```

<!-- grade_id: cell-650af470af428235, points: 1, grade: true, solution: false -->
```python
# Witdhraw more than balance:
my_current.withdraw(220)

assert my_current.balance == 0, "The created current account must not authorize withdraws when the amount is higher than the account balance"
```

<!-- grade_id: cell-8f5b292e2c0a7fc0, points: 1, grade: true, solution: false -->
```python
# Making a new deposit and checking balance:
my_current.deposit(750)
my_current.check_balance()

# New withdraw (should be blocked by limit):
my_current.withdraw(220)
assert my_current.balance == 750

# Charge current account annual fee:
my_current.apply_annual_fee()
print("\nChecking account balance after annual fee:")
my_current.check_balance()
assert my_current.balance == 730
```

#### End of the SLU8 Exercise Notebook.