# SLU08 Object Oriented Programming | Inheritance: Exercise Notebook

***

### Part 1
#### Exercise 1

Create a class named `Vehicle()` and define the `__init__()` method of the created class to have four arguments:

1. model
2. color
3. year
4. engine_type


Now, do the following:
- Assign these variables to attributes by using the self keyword.

- Define an extra attribute in the `__init__()` method, without user input, called `condition` and with default value of **"new"**.

- Create a `drive_it()` method that will override the vehicle condition to the string **"used"**

In the end, create the `my_car` object for a **1985** **silver** **"DeLorean"** with a **combustion** engine



<!-- grade_id: cell-152d925316906cc7, points: 0, grade: false, solution: true -->
```python
class Vehicle:
   
    ### BEGIN SOLUTION
    def __init__(self, model, color, year, engine_type):
        self.model = model
        self.color = color
        self.year = year
        self.engine_type = engine_type
        self.condition = "new"
    
    def drive_it(self):
        self.condition = "used"
        
my_car = Vehicle(model='DeLorean', color='silver', year=1985, engine_type='combustion')
    ### END SOLUTION
    
```

<!-- grade_id: cell-d6fb846dc8915dc9, points: 2, grade: true, solution: false -->
```python
assert my_car.model == 'DeLorean'
assert my_car.color == 'silver'
assert my_car.year == 1985
assert my_car.engine_type == 'combustion'
assert my_car.condition == 'new'

print("""If my calculations are correct, you are a proud owner of a brand %s, %s %s %s with a %s engine!
That's heavy!""" % (my_car.condition,
                    my_car.year,
                    my_car.color,
                    my_car.model,
                    my_car.engine_type))
print("---- all asserts passed ---- ")
```

Output:
```
If my calculations are correct, you are a proud owner of a brand new, 1985 silver DeLorean with a combustion engine!
That's heavy!
---- all asserts passed ----
```

<img src="./assets/delorean.jpg" width="400"/>

***

#### Exercise 2
Now just go and take you brand new car for a spin, using the `drive_it()` method of your `my_car` object.

<!-- grade_id: cell-49ead50591e8361d, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
my_car.drive_it()
### END SOLUTION

# after driving it for the first time, you vehicle condition will be updated:
print("Updated vehicle condition: " + my_car.condition)
```

Output:
```
Updated vehicle condition: used
```

<!-- grade_id: cell-c48e251cfd4b31e2, points: 1, grade: true, solution: false -->
```python
assert my_car.condition == 'used', "The condition of your car is not correct. Make sure you use the 'drive_it()' method"
print("I am sorry to inform you that the condition of your beloved DeLorean is now " + my_car.condition + "!")
print("---- all asserts passed ---- ")
```

Output:
```
I am sorry to inform you that the condition of your beloved DeLorean is now used!
---- all asserts passed ----
```

#### Exercise 3

As seen before, one of the benefits of classes is that we can create more complicated classes that inherit variables or methods from their **parent classes**. This saves us time and helps us build more complicated objects, since these **child classes** can also include additional variables or methods.
+ Create an `ElectricalVehicle` class that **inherits** from `Vehicle`;
+ Create a new method `get_battery_level()` that returns a (pseudo)random integer from 1 to 100 *
+ Then, create an electrical vehicle named `my_tesla`, for a **2020 black Model_S** with an engine type of **molten salt battery**

Make sure that `my_tesla` belongs to `ElectricalVehicle` Class.

***** use the `get_random_number()` function, already imported in the first line, for generating pseudo-random integers.

<!-- grade_id: cell-9d9a01b72fa64213, points: 0, grade: false, solution: true -->
```python
from utils import get_random_number

### BEGIN SOLUTION
class ElectricalVehicle(Vehicle):
    def get_battery_level(self):
        return get_random_number(1, 100)

my_tesla = ElectricalVehicle(model = 'Model_S', color = 'black', year = 2020 , engine_type = 'molten salt battery')
### END SOLUTION
```

<!-- grade_id: cell-7722214f55e01cee, points: 2, grade: true, solution: false -->
```python
# __init__() tests:
assert my_tesla.model == 'Model_S'
assert my_tesla.color == 'black'
assert my_tesla.year == 2020
assert my_tesla.engine_type == 'molten salt battery'
assert my_tesla.condition == 'new'

# battery level tests:
assert (my_tesla.get_battery_level() >= 1) & (my_tesla.get_battery_level() <= 100)
print("Battery levels(%):", (my_tesla.get_battery_level()))

# isistance tests:
assert (isinstance(my_tesla, ElectricalVehicle)) == True, 'my_tesla must belong to the ElectricalVehicle Child Class'
assert (isinstance(my_tesla, Vehicle)) == True, 'my_tesla must belong to the Vehicle Parent Class'
assert (isinstance(my_car, Vehicle)) == True
assert (isinstance(my_car, ElectricalVehicle)) == False

print("""\nWoah, Nikolas would be proud!
A %s, %s %s %s with a %s!
Again...That's heavy!""" % (my_tesla.condition,
                    my_tesla.year,
                    my_tesla.color,
                    my_tesla.model,
                    my_tesla.engine_type))
print('')
print(isinstance(my_tesla, ElectricalVehicle),'my_tesla belongs to the ElectricalVehicle Child Class.')
print(isinstance(my_tesla, Vehicle),'my_tesla belongs to the Vehicle Parent Class.')
print("---- all asserts passed ---- ")
```

Output:
```
Battery levels(%): 72

Woah, Nikolas would be proud!
A new, 2020 black Model_S with a molten salt battery!
Again...That's heavy!

True my_tesla belongs to the ElectricalVehicle Child Class.
True my_tesla belongs to the Vehicle Parent Class.
---- all asserts passed ----
```

<img src="./assets/tesla_edison_jk.jpg" width="400"/>

***

### Part 2
#### Exercise 4
Start by creating a new class `Rectangle` with the following:
+ Attributes:
    - **length**
    - **width**
+ Methods:
    - `get_area()`: returns the product of the rectangle **length** with **width**
    - `get_perimeter()`: returns the sum of double the **lenght** and double the **width**

<!-- grade_id: cell-e1f05519ef39b0f6, points: 0, grade: false, solution: true -->
```python
# class Rectangle:
### BEGIN SOLUTION
class Rectangle:
    def __init__(self, length, width):
        self.length = length
        self.width = width

    def get_area(self):
        return self.length * self.width

    def get_perimeter(self):
        return 2 * self.length + 2 * self.width
### END SOLUTION
```

<!-- grade_id: cell-c0da80882a9d38bc, points: 2, grade: true, solution: false -->
```python
# considering a 10,50 m by 20,50 m rectangle:
r = Rectangle(length=10.50, width=20.50)
assert r.get_area() == 215.25, 'Area computation is not correct!'
assert r.get_perimeter() == 62.0, 'Perimeter computation is not correct'
print('Congratulations, your calculations are correct!')
print("---- all asserts passed ---- ")
```

Output:
```
Congratulations, your calculations are correct!
---- all asserts passed ----
```

Feeling of the day:
<img src="assets/calculus.jpg" width="400"/>

#### Exercise 5
Create a `Square` class, as a child of the `Rectangle` class. Change only the `__init__()` method in a way that now it only accepts one argument 'side'. You can call the init method from the parent class and assign both length and width to be equal to 'side'

(hint: use the `super()` method - `Learning Notebook: 2.3 | The super() function`)

<!-- grade_id: cell-9262f02e916a66f1, points: 0, grade: false, solution: true -->
```python
# class Square
### BEGIN SOLUTION
class Square(Rectangle):
    def __init__(self, side):
        super().__init__(length=side, width=side)
### END SOLUTION
```

<!-- grade_id: cell-0d959343e7bd5637, points: 1, grade: true, solution: false -->
```python
# considering a 10,50 m by 10,50 m square:
s = Square(side=10.50)

assert s.get_area() == 110.25, 'Area computation is not correct!'
assert s.get_perimeter() == 42.0, 'Perimeter computation is not correct'

# Assert that the super() is implemented:
import inspect
error_message = 'Did you use the super approach?'
assert 'self.width' not in inspect.getsource(Square.__init__), error_message
assert 'self.length' not in inspect.getsource(Square.__init__), error_message
assert "super()." in inspect.getsource(Square.__init__), error_message

# Assert that methods are not re-implemented:
error_message =  'Do not re-implement the methods, use the ones from the Rectangle :)'
assert 'self.length' in inspect.getsource(Square.get_area), error_message
assert 'self.width' in inspect.getsource(Square.get_area),  error_message
assert 'self.length' in inspect.getsource(Square.get_perimeter), error_message
assert 'self.width' in inspect.getsource(Square.get_perimeter), error_message

print('Yes! Your square area (%s m2) and perimeter (%s m) computations are correct!' %(s.get_area(), s.get_perimeter()))
print("---- all asserts passed ---- ")
```

Output:
```
Yes! Your square area (110.25 m2) and perimeter (42.0 m) computations are correct!
---- all asserts passed ----
```

#### Exercise 6

Create a new class `Triangle`, independent from the other classes:
+ Attributes:
    - **base**
    - **height**
+ Methods:
    - `get_tri_area()` which calculates the area of the triangle (do you remember the formula? :D)

<!-- grade_id: cell-361d990792a542f3, points: 0, grade: false, solution: true -->
```python
# class Triangle:
### BEGIN SOLUTION
class Triangle:
    def __init__(self, base, height):
        self.base = base
        self.height = height

    def get_tri_area(self):
        return 0.5 * self.base * self.height
### END SOLUTION
```

<!-- grade_id: cell-a0f3c54e307aeee9, points: 1, grade: true, solution: false -->
```python
triangle = Triangle(base=10, height=20)

assert triangle.get_tri_area() == 100, 'The area of the triangle is miscalculated!'

print("The area, in cm2, of our triangle is", triangle.get_tri_area())
print("---- all asserts passed ---- ")
```

Output:
```
The area, in cm2, of our triangle is 100.0
---- all asserts passed ----
```

#### Exercise 7

Now for the `RightPyramid` class, create methods that compute the area of this geometric solid using both the `Square` and `Triangle` classes as parents.

Override the `get_area()` method in order to compute the equation of the RightPyramid area. **Hint:** use the `super()` function to  make everything easier.

So, you're in luck, geometry is not the subject for this SLU! Here are the formulas that you need:

* The area of a Right Pyramid is equal to the sum of the base_area with the its lateral_area.
* The Lateral area can be computed as the next image shows (1/2 * perimeter * slant_height)

<img src="assets/pyramid_lateral_area.jpg" width="500"/>

Other options do exist to compute the lateral area, like using the triangle_area. Feel free to use what is easier for your as this is about python, not geometry!

+ Attributes:
    - **square_side** (this is equivalent to one side of the square base, just like the previous exercises. So, base is also equal to 1/4 of the perimeter, p,  of the square base of the pyramid)
    - **slant_height**
+ Methods:
    - `get_area()`

<!-- grade_id: cell-18e78aa00977daf8, points: 0, grade: false, solution: true -->
```python
# class RightPyramid...
### BEGIN SOLUTION
class RightPyramid(Square, Triangle):
    def __init__(self, square_side, slant_height):
        self.square_side = square_side
        self.slant_height = slant_height
        self.height = slant_height
        self.length = square_side
        super().__init__(self.square_side)

    def get_area(self):
        base_area = super().get_area()
        perimeter = super().get_perimeter()
        return 0.5 * perimeter * self.slant_height + base_area

# alternative solution:    
#     def get_area(self):
#         base_area = super().get_area()
#         triangle_area = super().get_tri_area()
#         return triangle_area * 4 + base_area    
### END SOLUTION
```

<!-- grade_id: cell-7a3b92fb1669680d, points: 3, grade: true, solution: false -->
```python
Giza = RightPyramid(square_side=10, slant_height=20)

print("The area of our little Giza pyramid is: ", Giza.get_area())

# Assert instance:
assert isinstance(Giza, RightPyramid)
assert isinstance(Giza, Square), 'Probably your class is not properly defined, should be a child class of the Square classe'
assert isinstance(Giza, Triangle), 'Probably your class is not properly defined, should be a child class of the Triangle classe'

# Assert that the super() is implemented:
error_message = 'Did you use the super approach?'
assert 'super()' in inspect.getsource(Giza.__init__), error_message
assert 'super()' in inspect.getsource(Giza.get_area), error_message

# Assert calculations:
assert Giza.get_area() == 500, 'The area of the right pyramid is miscalculated!'
print("---- all asserts passed ---- ")
```

Output:
```
The area of our little Giza pyramid is:  500.0
---- all asserts passed ----
```

***

### Part 3

In the following exercice the goal it is to create a BankAccount program, composed by one parent class and two child classes, in order to manage the clients accounts, either through the main (current) account or the savings account.

This is the general idea:

<img src="assets/bank_account.jpg" width="400"/>



#### Exercise 8

First things first, the `BankAccount` class:

+ Attributes:
    - **customer**
    - **account_number**
    - **balance** with default value of 0 (zero) but passed as an argument anyway (you'll need it for Ex.9 and 10)
    - **currency = '€'** currency variable is always € but no need to pass it as an argument
+ Methods:
    - `deposit()`: 
        * Deposit **amount** into the bank account. (increase the balance). 
        * If the **amount** is invalid (< 0) it should return an alert message to the user
        * **Warning**: it should not deduce the amount and return (not print) a kind of `Invalid withdrawal amount` message.
    - `withdraw()`:
        * Withdraw **amount** from the bank account. (decrease the balance).
        * This method must ensure that the account has sufficient funds for the asked **amount**. 
        * **Warning**: Again, it should be > 0 and also return (not print) an alert if the value is less than zero or if it's bigger than the balance
    - `check_balance()`: print a statement of the account balance, example _"The balance of account number 48522888 is €1000"_

<!-- grade_id: cell-aec63d5088ce5f30, points: 0, grade: false, solution: true -->
```python
class BankAccount:
    """
    Abstract base class representing a bank account
    """
    ### BEGIN SOLUTION
    def __init__(self, customer, account_number, balance=0):
        """
        Initialize the BankAccount class with a customer, account number
        and opening balance (default value of 0)
        """
        self.customer = customer
        self.account_number = account_number
        self.balance = balance
        self.currency = '€'
        
    def deposit(self, amount):
        """
        Deposit amount into the bank account
        """
        if amount > 0:
            self.balance += amount
        else:
            return 'Invalid deposit amount:' + str(amount)
            
    def withdraw(self, amount):
        """
        Withdraw amount from the bank account, ensuring that are sufficient
        funds
        """
        if amount > 0:
            if amount > self.balance:
                return 'Insufficient funds!'
            else:
                self.balance -= amount
        else:
            return 'Invalid withdrawal amount:' + str(amount)
            
    def check_balance(self):
        """
        Print a statement of the account balance
        """
        return 'The balance of account number ' + str(self.account_number) + ' is ' + str(self.currency) + ' ' + str(self.balance)
    ### END SOLUTION
```

<!-- grade_id: cell-5fb3aa9dc9bb6e4c, points: 2, grade: true, solution: false -->
```python
# Start by creating a new account:
my_account = BankAccount(customer='John Doe', account_number=21457288)


# assert the class:
assert isinstance(my_account, BankAccount) == True, 'Your class is not properly defined'

# assert the properties of the class:
assert my_account.account_number == 21457288, 'You have to define an "account_number" attribute for your class!'
assert my_account.customer == 'John Doe', 'You have to define a "customer" attribute for your class!'
assert my_account.balance == 0, 'The default value of "balance" should be 0!'
assert 'balance=0' in inspect.getsource(my_account.__init__), 'balance=0 must be passed as an argument anyway'
assert my_account.currency == '€', 'Currency symbol is not working... Did you change the existing line?'

print("Nice, the account for %s was successfully created. His account number is %s and for now is account is empty. \n"
     % (my_account.customer, my_account.account_number))

# Lets deposit that big, fat check !
my_account.deposit(amount=25000)

assert my_account.balance == 25000, 'Your balance, after making a deposit should change!'

print("%s just made a deposit of €25'000. His new account balance is: %s%s \n"
      %(my_account.customer, my_account.currency, my_account.balance) )

# And then just go shopping (online... better be online nowadays):
my_account.withdraw(amount=2250)

assert my_account.balance == 22750, 'Your balance, after making a withdrawal should change or it is not correct!'
print("--(1/2)-- all asserts passed ---- ")
```

Output:
```
Nice, the account for John Doe was successfully created. His account number is 21457288 and for now is account is empty. 

John Doe just made a deposit of €25'000. His new account balance is: €25000 

--(1/2)-- all asserts passed ----
```

<!-- grade_id: cell-aaeb81da6d6cc191, points: 1, grade: true, solution: false -->
```python
# assert that a 'str' message is returned for invalid deposit/withdraw:
assert type(my_account.deposit(-10)) == str, 'Are you returning a string message for invalid deposit/withdraw'
assert type(my_account.withdraw(-10)) == str, 'Are you returning a string message for invalid deposit/withdraw'

# Finally let's ensure that your program does not accept invalid operations:
my_account.deposit(amount=-10)
assert my_account.balance == 22750, 'Your balance, after making an invalid deposit should not change!'
my_account.withdraw(amount=-10)
assert my_account.balance == 22750, 'Your balance, after making an invalid withdraw should not change!'

print("--(2/2)-- all asserts passed ---- ")
```

Output:
```
--(2/2)-- all asserts passed ----
```

#### Exercise 9

Next we need to develop the `SavingsAccount` class, as a child class for `BankAccount`:

+ Attributes:
    - *interest_rate*: default it to 0, must be passed in the `__init()` argument.
    - get the BankAccount attributes:
        * **hint1**: being super explicit here, the `__init__()` should have the following arguments - `customer, account_number, balance, interest_rate` (don't forget the required default values for some of them).
        * **hint2**: the position where you get the BankAccount attributes is important (should be the last thing you do in the `__init__()` of the SavingsAccount.
    
+ Methods:
    - `add_interest()`: add interest to the account at the rate *interest_rate*, thus updating the account balance.

<!-- grade_id: cell-8f1a5436b1aa5fd0, points: 0, grade: false, solution: true -->
```python
# class SavingsAccount...

### BEGIN SOLUTION
class SavingsAccount(BankAccount):
    """
    A class representing a savings account
    """
    def __init__(self, customer, account_number, balance=0, interest_rate=0):
        """
        Initialize the savings account
        """
        self.interest_rate = interest_rate
        super().__init__(customer, account_number, balance)
    
    def add_interest(self):
        """
        Add interest to the account at the rate self.interest_rate
        """
        self.balance *= (1. + self.interest_rate)
### END SOLUTION
```

<!-- grade_id: cell-2e5fd33c1f193d5a, points: 2, grade: true, solution: false -->
```python
# Creating a new savings account:
my_savings = SavingsAccount(customer='Jane Doe', account_number=41522887, interest_rate=0.0055, balance=10000)

# assert the class:
assert isinstance(my_savings, SavingsAccount) == True, 'Your class is not properly defined'
assert isinstance(my_savings, BankAccount) == True, 'Your class has no relation with the BankAccount class'

# assert the properties of the class:
assert my_savings.account_number == 41522887, 'You have to define an "account_number" attribute for your class!'
assert my_savings.customer == 'Jane Doe', 'You have to define a "customer" attribute for your class!'
assert my_savings.balance == 10000, 'The default value of "balance" should be 10000!'
assert my_savings.currency == '€', 'Currency symbol is not working... Did you changed the existing line?'

print("Nice, the account for %s was successfully created. Her account number is %s. \n"
     % (my_savings.customer, my_savings.account_number))

# Check account balance:
print("Checking account balance before interest:")
print(my_savings.check_balance())

# adding some interest to the existing account:
my_savings.add_interest()
print("\nChecking account balance after interest:")
print(my_savings.check_balance())

assert my_savings.balance == 10055, "The created savings account is not generating the correct interest (asserted with 5%, 0.0055) or any interest at all."
print("---- all asserts passed ---- ")
```

Output:
```
Nice, the account for Jane Doe was successfully created. Her account number is 41522887. 

Checking account balance before interest:
The balance of account number 41522887 is € 10000

Checking account balance after interest:
The balance of account number 41522887 is € 10055.0
---- all asserts passed ----
```

#### Exercise 10

Finally we need to implement the `CurrentAccount` class, as a child class for `BankAccount`, that has a similar structure of the previous exercise:

+ Attributes:
    - **annual_fee**
    - **transaction_limit**
    - All the BankAccount attributes
    
+ Methods:
    - `withdraw()`: override the parent class method and create an if clause with an extra condition, in order to limit the possible amount that a person can withdraw from a Current Account daily. **Hint:** use the *transaction_limit* variable you just initialized.
    - `apply_annual_fee()`: Deduct the annual fee from the account balance, making sure that the account never goes negative! **Hint:** use python's `max()` function

<!-- grade_id: cell-6d21acc19869fec0, points: 0, grade: false, solution: true -->
```python
# class CurrentAccount...

### BEGIN SOLUTION
class CurrentAccount(BankAccount):
    """
    A class representing a current (checking) account
    """
    def __init__(self, customer, account_number, annual_fee,
                transaction_limit, balance=0):
        """
        Initialize current account
        """
        self.annual_fee = annual_fee
        self.transaction_limit = transaction_limit
        super().__init__(customer, account_number, balance)
        
    def withdraw(self, amount):
        """
        Withdraw amount if sufficient funds exist and amount is less
        than the single tranasaction limit
        """
        if amount > self.balance:
            return 'Insufficient funds'
        
        elif amount > self.transaction_limit:
            return 'Exceeds the single transaction limit!'

        elif amount < 0:
            return 'Invalid withdrawal amount:' + str(amount)
            
        else: self.balance -= amount  
    
    def apply_annual_fee(self):
        """
        Deduct the annual fee from the account balance
        """
        
        self.balance = max(0., self.balance - self.annual_fee)
        
### END SOLUTION
```

<!-- grade_id: cell-4f6963493a0e8fe9, points: 1, grade: true, solution: false -->
```python
# create a new current account:
my_current = CurrentAccount(customer='Richard Roe', account_number=78300991, annual_fee=20., transaction_limit=200.)

# assert the class:
assert isinstance(my_current, CurrentAccount) == True, 'Your class is not properly defined'
assert isinstance(my_current, BankAccount) == True, 'Your class has no relation with the BankAccount class'

# assert the properties of the class:
assert my_current.account_number == 78300991, 'You have to define an "account_number" attribute for your class!'
assert my_current.customer == 'Richard Roe', 'You have to define a "customer" attribute for your class!'
assert my_current.balance == 0, 'The default value of "balance" should be 0!'
assert my_current.currency == '€', 'Currency symbol is not working... Did you changed the existing line?'

print("Nice, the account for %s was successfully created. Her account number is %s and for now is account is empty.\n"
     % (my_current.customer, my_current.account_number))
print("--(1/3)-- all asserts passed ---- ")
```

Output:
```
Nice, the account for Richard Roe was successfully created. Her account number is 78300991 and for now is account is empty.

--(1/3)-- all asserts passed ----
```

<!-- grade_id: cell-650af470af428235, points: 1, grade: true, solution: false -->
```python
# Witdhraw more than balance:
my_current.withdraw(amount=220)

assert my_current.balance == 0, "The created current account must not authorize withdraws when the amount is higher than the account balance"
print("--(2/3)-- all asserts passed ---- ")
```

Output:
```
--(2/3)-- all asserts passed ----
```

<!-- grade_id: cell-8f5b292e2c0a7fc0, points: 1, grade: true, solution: false -->
```python
# Making a new deposit and checking balance:
my_current.deposit(amount=750)
my_current.check_balance()

# New withdraw (should be blocked by limit):
my_current.withdraw(amount=220)
assert my_current.balance == 750

# Charge current account annual fee:
my_current.apply_annual_fee()
print("\nChecking account balance after annual fee:")
my_current.check_balance()
assert my_current.balance == 730
print("--(3/3)-- all asserts passed ---- ")
```

Output:
```
Checking account balance after annual fee:
--(3/3)-- all asserts passed ----
```

#### End of the SLU08 Exercise Notebook.