# SLU11 Command Lines & Text Editor - Exercise Notebook
***

```python
# For evaluation purposes, please import this package:
import hashlib
```

### Exercise 1: Linux Command Line Interpreter
What is the name of the command line interpreter in LINUX?

- Option A: PowerShell
- Option B: UNIX Shell
- Option C: Bourne Again Shell (BASH)

<!-- grade_id: cell-e435dad275e6d3e8, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the option value of the option that you think is correct to the variable exercise_1.
# For example, if you think that option A is correct, do: exercise_1 = "A"

# exercise_1 =
### BEGIN SOLUTION
exercise_1 = "C"
### END SOLUTION
```

<!-- grade_id: cell-477ba88499d4f59b, points: 2, grade: true, solution: false -->
```python
ex_1 = '6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d'
assert isinstance(ex_1,str), "The variable should be a string."
assert ex_1 == hashlib.sha256(bytes(exercise_1, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 2: Hierarchical directory structure
How are Linux files organized in terms of structure?

- Option A: Folders and sub-folders
- Option B: UNIX Shell
- Option C: Directories and sub-directories

<!-- grade_id: cell-9bb39a7a94cf8caf, points: 0, grade: false, solution: true -->
```python
# exercise_2 =
### BEGIN SOLUTION
exercise_2 = "C"
### END SOLUTION
```

<!-- grade_id: cell-f82b2df42e06ea99, points: 2, grade: true, solution: false -->
```python
ex_2 = '32096c2e0eff33d844ee6d675407ace18289357d'
assert isinstance(ex_2,str), "The variable should be a string."
assert ex_2 == hashlib.sha1(bytes(exercise_2, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 3: Where is home?
Imagine that you have a Linux user called `tortoise`, what is the correct path for this user's *home directory*?

- Option A: */root*
- Option B: */tortoise*
- Option C: */root/tortoise*
- Option D: */home/tortoise*

<!-- grade_id: cell-a894a58ceb39eaa2, points: 0, grade: false, solution: true -->
```python
# exercise_3 =
### BEGIN SOLUTION
exercise_3 = "D"
### END SOLUTION
```

<!-- grade_id: cell-5a2440598289b1e3, points: 3, grade: true, solution: false -->
```python
ex_3 = '50c9e8d5fc98727b4bbc93cf5d64a68db647f04f'
assert isinstance(ex_3,str), "The variable should be a string."
assert ex_3 == hashlib.sha1(bytes(exercise_3, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 4: Changing directory
Now, for our same `tortoise` user, if our current working path is `/var/lib/systemd/`, how do we go this users *home directory* ?

Option A: pwd
Option B: cd /home
Option C: cd
Option D: cd tortoise

<!-- grade_id: cell-29db176b41ae1c07, points: 0, grade: false, solution: true -->
```python
# exercise_4 =
### BEGIN SOLUTION
exercise_4 = "C"
### END SOLUTION
```

<!-- grade_id: cell-e4ea53657d321c19, points: 3, grade: true, solution: false -->
```python
ex_4 = '484d52691fcadbfabec5a318d1cf9692c7f293cbc8c1d5f22b2d839b'
assert isinstance(ex_4,str), "The variable should be a string."
assert ex_4 == hashlib.sha224(bytes(exercise_4, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 5: New directory
We need to create a new working directory for `tortoise` files in this current working directory, let's call it `/dev`, what is the correct command we should use?

- Option A: *cat /dev*
- Option B: *mkdir dev*
- Option C: *make /dev*
- Option D: *touch dev*

<!-- grade_id: cell-befccab89f6c99df, points: 0, grade: false, solution: true -->
```python
# exercise_5 =
### BEGIN SOLUTION
exercise_5 = "B"
### END SOLUTION
```

<!-- grade_id: cell-8e25f4730d887ad4, points: 3, grade: true, solution: false -->
```python
ex_5 = 'd56cc1ac775bdf54c4536f6ef08de810454487513640e201af923c79'
assert isinstance(ex_5,str), "The variable should be a string."
assert ex_5 == hashlib.sha224(bytes(exercise_5, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 6: New file
`tortoise` is not happy. This user needs a new empty file where is going to store some logs, called `dump.txt`. 

What command should we use?

- Option A: *cat /dump.txt*
- Option B: *rmdir dump.txt*
- Option C: *ls dump.txt*
- Option D: *touch dump.txt*

<!-- grade_id: cell-d629c9437fa027f7, points: 0, grade: false, solution: true -->
```python
# exercise_6 =
### BEGIN SOLUTION
exercise_6 = "D"
### END SOLUTION
```

<!-- grade_id: cell-9b757ee86b6b31a4, points: 3, grade: true, solution: false -->
```python
ex_6 = '7497e88b8c64062edfb2f30b9c5f845552b841d1fbae97ab3f224f74'
assert isinstance(ex_6,str), "The variable should be a string."
assert ex_6 == hashlib.sha224(bytes(exercise_6, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 7:
Now that `tortoise` stored every log entry that he could in the `dump.txt` file, he needs to check what is stored in the file in a simple and easy way.

What is the correct command for that procedure?

- Option A: *cat dump.txt*
- Option B: *rmdir dump.txt*
- Option C: *ls dump.txt*
- Option D: *touch dump.txt*

<!-- grade_id: cell-1792bbe6ebc8564d, points: 0, grade: false, solution: true -->
```python
# exercise_7 =
### BEGIN SOLUTION
exercise_7 = "A"
### END SOLUTION
```

<!-- grade_id: cell-d57ae086b71ffed4, points: 3, grade: true, solution: false -->
```python
ex_7 = 'ad14aaf25020bef2fd4e3eb5ec0c50272cdfd66074b0ed037c9a11254321aac0729985374beeaa5b80a504d048be1864'
assert isinstance(ex_7,str), "The variable should be a string."
assert ex_7 == hashlib.sha384(bytes(exercise_7, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 8: Visual Studio Code

Managing to install and set VSCode up is part of this SLU, so if you have any problems please do contact an instructor! 

It is very important that you get it running properly.

- Option A: Rock On! Everything is working fine :)
- Option B: NO, I am still struggling with it. It's this week's SLU instructor fault!!!

<!-- grade_id: cell-82043c053fb1c39d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
exercise_8 = "A"
### END SOLUTION
# exercise_8 =
```

<!-- grade_id: cell-fa61be6aa4097dad, points: 1, grade: true, solution: false -->
```python
ex_8 = '559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd'
assert isinstance(ex_1,str), "The variable should be a string."
assert ex_8 == hashlib.sha256(bytes(exercise_8, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

### Exercise 9: BASH and VSCode

To finish this SLU in style, try to perform the following task using VSCode and the BASH command line interface that comes with it:

Remember that for the Linux user **tortoise**, `./` is translated to `/home/tortoise`.

1. In your home directory create the following directory structure:
    - ./my_folder_1/my_sub_folder_1/
    - ./my_folder_2/my_sub_folder_2/
    
    
2. Navigate to the **my_sub_folder_1** and create a new Python file called **my_code.py** with the `touch` command.


3. With VSCode editing abilities, paste the following snippet inside the **my_code.py** file and save it:

```python
import os
os.getcwd()
```

Output:
```
'/home/nunohelibeires/ds-prep-course-instructors/Week 7/SLU11 - Command Lines & Text Editor'
```

4. Use the `cat` command to confirm everything is correct inside the python file.


5. Run the **my_code.py** file using the BASH command line.
    - *Expected result: ./my_folder_1/my_sub_folder_1/*
    
    
6. Use the `cp` command to copy the python file to ./my_folder_2/my_sub_folder_2/


7. Run the **my_code.py** file using the BASH command line.
    - *Expected result: ./my_folder_2/my_sub_folder_2/*