# SLU13 - Git Intermediate - Exercise notebook

## Start by importing these packages

```python
# for evaluation purposes
import hashlib
def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()
```

# Exercise 1: Git remotes


What is the main purpose of a remote repository?
- A: Committing changes to your files.
- B: Sharing work with your teammates.
- C: To serve as a backup to your local files.


<!-- grade_id: cell-acb1a32dbeb026ee, points: 0, grade: false, solution: true -->
```python
# Uncomment the correct answer
# exercise_1 = "A"
# exercise_1 = "B"
# exercise_1 = "C"

# YOUR CODE HERE
raise NotImplementedError()
# exercise_1 =
```

<!-- grade_id: cell-126db698519af9eb, points: 0, grade: true, solution: false -->
```python
assert _hash(exercise_1) == '90a8cbb480', (
    'Not quite right! Review Section 1 of the Learning Notebook.') 

print("---- all asserts passed ---- ")
```

# Exercise 2: Git remotes

Did your review how to create and clone a repository [here](https://github.com/LDSSA/ds-prep-course#setup-git-and-github)? Good!

Go to [Github](https://github.com/), and create a new repository called *SLU13* and then clone it to your machine.

It's the same thing you did for setting up the *ds-prep-workspace* repository, but with a different name:
1. You don't need to setup git and Github again, that part is already covered
2. Follow the instruction on [Setup your Workspace Repository](https://github.com/LDSSA/ds-prep-course#setup-your-workspace-repository), but name your repository *SLU13*
3. Follow the instruction on [Cloning the workspace](https://github.com/LDSSA/ds-prep-course#cloning-the-workspace) for cloning the remote repository into your machine.

That's it for this exercise. We will use this new repository for the remaining exercises

# Exercise 3: Pushing changes into remote repository

Do the following workflow:
1. On your local *SLU13* repository, create a groceries.txt file and add some grocery items to it
    1. you can add more files if you want to
2. Stage and commit your changes
3. Do some other change to your files
4. Stage and commit your changes
5. Verify, by using `git status` and `git log`, that you have 2 commits in the local repository that are not in the remote repository
6. Push the changes to the remote repository
7. Verify, by using `git status` and `git log`, that your local and remote repositories are in sync.

You don't have to follow the exact workflow above. Feel free to experiment on this one :) The idea is to get used to the commands. You will use them a lot, so start getting that muscle memory now.

 # Exercise 4. Pulling changes from remote
 
 Try to do the same workflow we did on the learning notebook:
 1. Use your browser to go to your *SLU* repository on Github (probably `https://github.com/<your_github_userid_here>/SLU13`)
 2. Edit your groceries.txt file
 3. Commit your changes
 4. Check that `git status` and `git log` don't show any differences. If you feel like it, disconnect your computer from the network and check that those commands keep working :)
 5. Try using `git remote` to check if the repositories are synched
 6. Pull the changes you did in the remote into your local repository 

 # Exercise 5. Stash Changes
 
 Try the following workflow:
 1. Create a new file named *books.txt*
 2. Add a few books you'd like to read during the quarantine
 3. Stash the changes (hint: remember that stash does not save untracked files)
 4. Add a new item to the *groceries.txt* file
 5. Commit those changes
 6. Pop the changes you have in the stash
 7. Add some more books if you want to and, when you're ready, commit
 6. Check, using whatever method you like, that the local repository is 2 commits ahead of the remote repository
 7. Push the commits to the remote repository

 # Exercise 6. Create and push branches
 
 Try the following workflow:
 
 1. Create a new branch based on **master**
 2. In the new branch create a new file **foods_to_try.txt** and add some content to it
 3. Commit the changes and push the branch to the remote
 4. Check on github that the branch is there

# Exercise 7. Merging branches

I'm leaving out a lot of step by step instructions in this exercise and the next one. You should by now understand what you have to do. Try the following:

1. Merge the branch you created on exercise 6 into *master*. Verify that it was a fast-forward merge.
2. Do a 3-way merge of a branch into master. This might be a tough one. Go through the learning notebook carefully. Everything you need is in there.
3. Push *master* to the remote.
4. (Bonus): Do a google search on how to delete a branch and delete the branches you created, both on the local repository and the remote.

# Exercise 8. Merge conflicts
One-liner for this one:

- Create a merge conflict and solve it any way you like. Extra points if you solve it using several methods. Tip: you can use `git merge --abort` before committing and use the same conflict again and again.
