# SLU01 - Exercise notebook

## Start by importing this package

```python
# Just for evaluating the results.
import math
import hashlib
```

## Exercise 1: Data Types

## 1.1) Scientific notation

<!-- grade_id: cell-a5bd49d98abf2410, points: 0, grade: false, solution: true -->
```python
# Calculate the distance (in meters) a light particle travels for 12 seconds.

# Assign the speed of light (2.99*10^8 m/s) to the variable speed_light.
# Use the scientific notation and 2 decimal places.

# speed_light = ...
# YOUR CODE HERE
raise NotImplementedError()


# Assign the distance travelled by the particle to variable distance_travelled.

# distance_travelled = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6b465d1d500666a0, points: 2, grade: true, solution: false -->
```python
assert isinstance(speed_light, float), "Review the data type of speed_light."
assert math.isclose(speed_light,2.99E8,abs_tol=1E5), "Check your speed!"
assert isinstance(distance_travelled, float), "Review the data type of distance_travelled."
assert math.isclose(distance_travelled,2.99E8 * 12,abs_tol=1E5), "Did you use the right time?"
```

## 1.2) Data type conversion

The data types of several variables got mixed up. Convert them to the target type.

**Do not change** the values of the original variables.

### 1.2.1) To string

<!-- grade_id: cell-cf0b9cd761180962, points: 0, grade: false, solution: true -->
```python

# Assign to variable size_string the value of variable size converted into a string.
size = 1.78

# size_string = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ac16bba863ed00af, points: 2, grade: true, solution: false -->
```python
size_string_hash = '5d48c7c6dce082d397fecf99b8f1ac7fa7adf02bd38fcbb1f7f9f13fba7a2d45'

assert isinstance(size_string, str), "Review the data type of size_string."
assert size_string_hash == hashlib.sha256(bytes(size_string, encoding='utf8')).hexdigest(), "The value of size_string is incorrect."
```

### 1.2.2)  To integer

<!-- grade_id: cell-ca00a15a43742c1c, points: 0, grade: false, solution: true -->
```python
# Assign to variable age_int the value of variable age converted into a integer. 
age = "34"

# size_string = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-43f5849625865cc0, points: 2, grade: true, solution: false -->
```python
assert isinstance(age_int, int), "Review the data type of age_int."
assert math.isclose(age_int, 34, abs_tol=0.1), "The value of age_int is incorrect."
```

### 1.2.3)  To float

<!-- grade_id: cell-cce49b7314add616, points: 0, grade: false, solution: true -->
```python
# Assign to variable weight_float the value of variable weight converted into a float. 
weight = "78.4"

# weight_float = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-97305f11f776d1a5, points: 2, grade: true, solution: false -->
```python
assert isinstance(weight_float, float), "Review the data type of weight_float."
assert math.isclose(weight_float, 78.4, abs_tol=0.01), "The value of weight_float is incorrect."
```

# Exercise 2: Operations

## 2.1) Arithmetic operations

Solve the math problems and assign the results to the variables.

Use arithmetic operations to solve the questions.

If the answer has no decimal part, assign it as an **integer** if necessary. Otherwise leave it as a **float**.

### 2.1.1)  Basic arithmetic

<!-- grade_id: cell-68e58a91a765c6cb, points: 0, grade: false, solution: true -->
```python
# To study Data Science, John read 2 books. 
# He read the first one in one week with 25 pages everyday. 
# He read the second book in 12 days with 23 pages everyday. 
# What is the total number of pages that John read? 

# pages_read = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b09b6a3df43346e0, points: 2, grade: true, solution: false -->
```python
assert isinstance(pages_read, int), "Review the data type of pages_read."
assert math.isclose(pages_read, 451, abs_tol=0.001), "The value of pages_read is incorrect."
```

### 2.1.2) More basic arithmetic

<!-- grade_id: cell-0f95f2283fa7b5ed, points: 0, grade: false, solution: true -->
```python
# Steve used his savings of 2300 euros for some purchases.
# He spent 3/4 of his savings on furniture. 
# He then spent 1/2 of his remaining savings on a fridge. 
# How much money had Steve left on his savings after the purchases?

#savings_left =
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b15473e30437404d, points: 2, grade: true, solution: false -->
```python
assert isinstance(savings_left, float), "Review the data type of savings_left."
assert math.isclose(savings_left, 287.5, abs_tol=0.2), "The value of savings_left is incorrect."
```

### 2.1.3) Even more arithmetic

<!-- grade_id: cell-4a8b61ca7fd9597b, points: 0, grade: false, solution: true -->
```python
# John, Mary, Louis and Steve have 175 cards to use in a game. 
# They want to share them equally. 
# How many cards should each one take and how many cards are left on the deck?

# cards_per_player = ...
# cards_left = 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-80de177e3edd27bc, points: 2, grade: true, solution: false -->
```python
assert isinstance(cards_per_player, int), "Review the data type of cards_per_player."
assert math.isclose(cards_per_player, 43, abs_tol=0.001), "The value of cards_per_player is incorrect."
assert isinstance(cards_left, int), "Review the data type of cards_left."
assert math.isclose(cards_left, 3, abs_tol=0.001), "The value of cards_left is incorrect."
```

### 2.1.4) Shortcut operators

<!-- grade_id: cell-ca40d76b6b0fd084, points: 0, grade: false, solution: true -->
```python
# Louis has invested 2500 euros in a savings account. 
# The investment gives returns of 5% of the value invested every year.
# Louis reinvests all the returns in the same savings account.
# How much money is there in the savings account at the end of the 3rd year?


# TIP: you can assign as many auxiliary variables as you need. 
# Or change the value of a variable after the initial assignment.

# final_savings =
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-822b5ec7dfe02776, points: 2, grade: true, solution: false -->
```python
assert isinstance(final_savings, float), "Review the data type of final_savings."
assert math.isclose(final_savings, 2894.0625, abs_tol=0.1), "The value of final_savings is incorrect."
```

## 2.2) String operations

### 2.2.1) Concatenation

<!-- grade_id: cell-0a458560de5f3c4e, points: 0, grade: false, solution: true -->
```python
# Concatenate the joke with 3 times the response leaving a space between each string.
# The result should be: "How does Darth Vader like his toast? On the dark side. Ha Ha Ha"
joke = "How does Darth Vader like his toast? On the dark side."
response = "Ha"
# joke_reception = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-65f0dc204b4b0812, points: 2, grade: true, solution: false -->
```python
joke_reception_hash = '3b544b8a3e3c38076931e145008372476700160b6c106dea0aede521192a5a88'

assert isinstance(joke_reception, str), "Review the data type of joke_reception."
assert joke_reception_hash == hashlib.sha256(bytes(joke_reception, encoding='utf8')).hexdigest(), "The value of joke_reception is incorrect."
```

### 2.2.2) String formatting

<!-- grade_id: cell-d2b9abfbcaf35a41, points: 0, grade: false, solution: true -->
```python
# Use one of the string formatting methods to insert the values of the variables into a string.
# Write the string in a single statement but DO NOT write the quantities "by hand" and DO NOT use concatenations.
# The result should be: "In the fridge there are 12 yogurts, 0.5 liters of milk and 8 eggs."
yogurt = 12
milk = 0.5
eggs = 8

# inventory = "In the fridge (...)"

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9f9722689d8a954e, points: 2, grade: true, solution: false -->
```python
inventory_hash = 'd48f0e70729026e365046f03606e62d2dfab6b3b1102300e38efbeae63ce156d'

assert isinstance(inventory, str), "Review the data type of inventory."
assert inventory_hash == hashlib.sha256(bytes(inventory, encoding='utf8')).hexdigest(), "The value of inventory is incorrect."
```