# SLU 07 - Exercise notebook 

Before you start:
>- This notebook is divided into 8 questions: 
    - Question 1 to 6 are easy, practice questions, expect that part to take about 2 hours. 
    - Questions 7 and 8 are hard question, expect to spend at least another 3 hours on them. They are close to the level of difficulty to pass the Academy test. 
    
> - If you are unsure about your answer, run the asserts (the tests). Even if it is a bit incomplete, run the asserts, and see if they give you clues. 
    - In the last 2 questions the asserts will be a bit less helpful, so that you can get some practice of dealing with genuine python errors. 
    - Read the instructions carefully. Everything that is there is critical
    - Don't worry if these concepts feel weird at first. Classes are not easy, you are now getting into advanced stuff! 
    
> - If you are stuck, remember to:
    - Play around! Open a new cell, and mess around with the objects. Instantiate them and see if they complain. See what they print, what they return, etc. Remember, in real life there are no asserts to guide you!  
    - Always read the error messages carefully. They should be good clues about what is wrong. 
    - As always, don't be afraid to ask questions! 


----

### Exercise 1 - Playing around with objects 

This exercise is very simple. We are just going to make a couple new fruits, like we did in the first Learning Notebook. 

Make the following two fruits: 
- pears, which is called "pears", has 4 units, costs 3 euros per unit, and expires in 14 days
- tangerines, which is called "tangerines", has 2 units, costs 6 euros per unit, and expires in 2 days

This exercise is very easy, just look up how it was done in the Learning notebook and do a couple of them here. 

Reminder: the Fruit class has the following attributes: 
- `price_per_unit`
- `name`
- `nr_units`
- `days_until_expired`

> Protip: Note that you can use the tab auto-complete to help you.  
> Try writing `Fruit(pr` and hitting tab. It should auto-complete to price! 


<!-- grade_id: cell-4640dc55a4c658cd, points: 0, grade: false, solution: true -->
```python
from utils import Fruit

# pears = ...
# tangerines = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f4f479ecb4b7c3ef, points: 1, grade: true, solution: false -->
```python
from utils import exercise_1_grading
assert (pears and tangerines), 'Did you forget to make one of the fruits?'
exercise_1_grading(pears, tangerines)
print('Great success! Exercise passed!')
```

----
### Exercise 2 - add up the price of fruits 

Great! For your second exercise you will add the price of all the fruits: 
- apples
- bananas 
- oranges 
- pears
- tangerines 

The price is the product of the price per unit and number of units. But remember, the Fruit object already comes with a `calculate_price` method! 

Create a function called `calculate_price_of_all_fruits`, which takes a list called `all_fruits` that contains fruit objects and returns the total price. We've already filled in some things for you. 

Remember: the function should take any list of fruits, don't hard code any particular list! 

<!-- grade_id: cell-d5a600cefd3c41c7, points: 0, grade: false, solution: true -->
```python
def calculate_price_of_all_fruits(all_fruits): 
    """
    Calculates the price for every fruit in the all_fruits list, 
        and returns the total price of fruits in the list
    Args: 
        all_fruits (list): a list with fruit objects 
        
    Returns: 
        total_price (int): the total price of the fruits in the list
    """
    
    # hint: this may require a for loop, and a variable to keep track of the total price 

    # YOUR CODE HERE
    raise NotImplementedError()
    
# this should return 24 
print(calculate_price_of_all_fruits([pears, tangerines]))
```

<!-- grade_id: cell-ee1a997f801ac0bb, points: 1, grade: true, solution: false -->
```python
from utils import exercise_2_grading
exercise_2_grading(calculate_price_of_all_fruits)
print('Great success! Exercise passed!')
```

---
### Exercise 3 - Using objects with other objects 

In this exercise you will do something very similar to the learning notebook. You will simply instantiate a basket, create some toilet paper, and then will add and remove stuff until you have the following: 

- Total value of the basket should be between 30 and 35 
- the basket must have at least 3 types of fruit
- the basket must not contain any items that will expire in the next 8 days 
- the basket must have at least some toilet paper. It should be ... 
    - named "Soft but cheap" 
    - cost 4 euros
    - have "Double leaf". 
    - The expiration date and number of units are up to you, they won't be tested

Reminders: 
```
- You already have some instance variables to work with: 
    - apples 
    - oranges
    - bananas 
    - pears 
    - tangerines 
- You only have the Class variable Toiletpaper to work with. 
- You'll need to create some toilet paper before you can use it. 
    - Call the variable toilet_paper.
    - Give it nr_units = 1 
```

_Note: obviously you can figure out exactly what you need to add ahead of time without trial and error. However, it's good practice to just use the objects. Try adding, checking the price, removing, etc until you have an answer that meets all the conditions._

<!-- grade_id: cell-cffd2f4f2528de89, points: 0, grade: false, solution: true -->
```python
from utils import Toiletpaper, Basket, get_fruits
apples, bananas, oranges = get_fruits()  # you already have some fruits

# my_basket = ... 
# toilet_paper = ...
# my_basket.<do something> 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-277ef4e6f1a47d0d, points: 2, grade: true, solution: false -->
```python
from utils import exercise_3_grading 
exercise_3_grading(my_basket, toilet_paper)
print('Great success! Exercise passed!')
```

----
### Exercise 4: Making a basic object 

For this exercise, you will create a few `Student`s. 

Here is what you need to do: 
1. Create a `Student` class. Each student should be initialized with a `name`, an `age`, and an `average_grade`. You don't know however what that `average_grade` should be, so it should always be initialized as None. 

Then create 3 students: 
> `minh`, called "Minh Hoang", aged 15  
> `maria`, called  "Maria Dominguez", aged 13  
> `sam`, called "Sam Hopkins", aged 8  



<!-- grade_id: cell-2e5ef92f86968bc5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-61f3abe376b0ed6d, points: 2, grade: true, solution: false -->
```python
from utils import exercise_4_grading
assert Student
assert (minh and maria and sam), 'did you not forget to create someone?'
exercise_4_grading(Student, maria, minh, sam)
print('Great success! Section passed!')
```

---
### Exercise 5: working with methods 
For this exercise you will create a `School` class. 

Every school will be instantiated without any students, just an empty list (a bit like the basket, hint hint) called `students`. Each school should also have a `name`. 

It should also have a method called `accept_student`, which takes a student and appends them to the `students` list. 

<!-- grade_id: cell-a9722b9debaeb161, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a9853f7b4e57fe87, points: 2, grade: true, solution: false -->
```python
from utils import exercise_5_grading
assert School
test_school = School(name='Mean Girls High')
exercise_5_grading(School, test_school)
print('Great success! Section passed!')
```

---
### Exercise 6: calling methods with other methods 
We were a bit naive, and students who are too young to get into highschool are getting in!  

Create a `Highschool` class, that is the same as `School` (just copy paste what you wrote if you want), but has two extra methods:
- `consider_student_application`, which will take a student, check if the age is equal or more than 12 and, if it is, will use the `accept_student` method to accept the student.  
- `show_all_students`, which will loop over the students that have been accepted and print their names.

After you have created the school class, use it to create `happy_highschool` as an instance of `Highschool`, which is named `The happy highschool`. 

Finally, make this happy_highschool consider the application of `minh`, `maria` and `sam`, who you created in exercise 4. (you don't need to create them again).  Remember that they students all have the attributes `age` and `name`. 

<!-- grade_id: cell-ef56f4c9a022e15a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9eae48817b9d3b7f, points: 4, grade: true, solution: false -->
```python
from utils import exercise_6_grading
assert Highschool
assert happy_highschool
exercise_6_grading(Highschool, happy_highschool)
print('Great success! Section passed!')
```

----
> Before you continue:
- _The exercises up to here were good for practice. The last two are hard ones._ 
- The last 2 exercises have no new material. They will just have less hand-holding and be more demanding. If you get in trouble, see your own solutions above for clues! 
- You can of course ask for help at any time, but try suffering through these two on your own if you can, it's very good practice for the real world :)  
- _Expect to take quite a while (probably over 2 hours) to get exercise 7 and 8 right._
- _It's often useful to draft the "skeleton" of your code before jumping in. What are the steps? What are the methods of each class you will need? Sometimes writing on a piece of paper is a good support._ 
- _This section will test you on what you learned, but also includes functions, and other things you have learned before. You may need to check a few past SLUs for this one_
- _After you create your Employee class, make an employee called frank, and give him a salary. See if that works. The same goes for the Company class. It's crucial to gain the habit of testing your stuff as you go, to see if it is working._
    - _the asserts will give you hints as to where you have a problem, but it's your responsibility to create instances, find the actual error, and try to fix it!_ 
- _This is the last block in this SLU! If you got this, you are getting pretty good at objects, and are probably ready to deal with the Academy entry exercise that has this SLU's material._

---
### Exercise 7: the companies and the employees 

Your objective is to create some companies, and some employees. Each company will hire some employees and have the ability to give them all raises. You will also need to check what the mean salary of each company is.   


About the Employee class: 
- every employee should have a `name` and a `salary` (don't worry about the fact that they have a salary before they join the company, let's keep things simple).


About the Company class: 
- The company should have the following attributes:
    - `name`, which will be passed when instantiating companies 
    - `list_of_employees`, which should start empty when the company is instantiated. 
- The methods for `Company` should be as follows: 
    - `hire` --> ads an `employee` to the company's `list_of_employees`
    - `give_a_raise`, which has two arguments: 
        - `employee` 
        - `raise_fraction` takes number between 0 and 1. For instance, if an employee has a salary of 1000 and the raise_fraction is 0.1, they should get a 10% increase in their salary (to 1100). 
    - `get_mean_salary`, which calculates the mean salary of all company employees and returns it. 
    - `give_everyone_a_raise`, which takes an argument of `raise_fraction` (for instance 0.2) and gives that fraction raise to every employee (in this case a 20% raise for everyone). 

<!-- grade_id: cell-05a39eba6e7e96c1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e429403c7f5eff78, points: 4, grade: true, solution: false -->
```python
from utils import exercise_7_grading
assert Company, Employee
exercise_7_grading(Employee, Company)
print('Great success! Section passed!')
```

---
### Exercise 8: putting your classes to use! 
You will then put these classes to use. You will then create two companies, and 6 employees:

companies: 
    - `generous_company`, called "Generous company Ltd."
    - `stingy_company`, called "Mr. Burns Enterprises Ltd."
employees (we don't care what you name the variables):  
        - "Hugo C", salary 5
        - "Juliana R", salary 6
        - "Ines P", salary 5
        - "Pavel N", salary 7 
        - "Carol M", salary 8
        - "Bruno A", salary 6

The first 3 employees (Hugo, Juliana and Ines) will be hired by `generous_company`, and the last 3 will be hired by `stingy_company`. 

You will check the mean salary right after both companies have finished hiring. 

Each company will give all of their employees raises 5 times, but `generous_company` will give 20% raises, and `stingy_company` will give only 5% raises. 

You will then check the mean salary for each company again. 

Remember, for this exercise you don't need to create any new classes. You just need to create instances of companies and employees, using the classes you defined above, and use the methods you have already defined. You'll need to look at your answer from exercise 7 a few times to remember how your classes work. 

This is the last exercise, good luck! 

<!-- grade_id: cell-a43a4e57177eca9d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1d7bb58d538228da, points: 4, grade: true, solution: false -->
```python
from utils import exercise_8_grading
assert generous_company and stingy_company
exercise_8_grading(generous_company, stingy_company, Employee)
print('Great success! Everything passed, congratulations!!!')
```