# SLU14 - Modules and Packages 

#### Start by importing the following module

```python
#used for evaluation
import utils
```

### Q1 What is a module in python?

a) A set of attributes;   
b) A set of files with `.py` extention;   
c) A file with `.py` extention;   
d) A file with `.ipynb`extention;

<!-- grade_id: cell-ec299e14ecce1f27, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-60bb3a2487965c64, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q1, "Wrong answer."
```

### Q2 How should we organize code in a module?

a) In attributes: classes, functions and variables;      
b) With no special organization, but with all code inside `if __name__=="__main__"`;   
c) With no special organization at all;   
d) In attributes: classes, functions and variables. If we also want to run it as a script, the entry point should be inside  `if __name__=="__main__"`;

<!-- grade_id: cell-2f1a5dfccc2eee9b, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-18368bb83c4a41b4, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q2, "Wrong answer."
```

### Q3 Why modules should be organized that way? 

a) Just because it is better organized and easier to debug;   
b) In that way we can use the attributes outside of the file and still run it as a script.

<!-- grade_id: cell-7a2c49000d940397, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cd2427e89b7707cb, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q3, "Wrong answer."
```

### Q4 What are the main advantages of organizing our projects by modules? Pick the one that is less relevant. 

a) It enable us to separate code by features and avoid duplicated code;   
b) Easier to debug, mantain and scale;   
c) Allow us to run modules as scripts;   
d) Can be packaged together and used in other projects;  

<!-- grade_id: cell-a4658c8feaa4f078, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9a98124ab22debe2, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q4, "Wrong answer."
```

### Q5 What we should write inside our module, right before the entry point that will be executed when we run it as a script?

a) Nothing;   
b) `if __name__=="__main__"`

<!-- grade_id: cell-aea5631b25c21caa, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-99cbfb671d0969b3, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q5, "Wrong answer."
```

### Q6 Inside directory shapes, create a module named `triangle`  with the following attributes:

- A `class` named `Triangle` with two inputed variables, `base` and `height`, that are attributes of the class. This class should have also a method names `get_area`that calculates the area in the following way: `(base*height)/2`.    
- A function names `is_equilateral` that accepts three variables as an input. If they have all the same value, return True, False otherwise.   
- A variable names `description` with the following string value `this is a module named triangle`. 

<!-- grade_id: cell-a8c79db3d36a242b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Q7 Import the class `Book` , that belongs to module `book`, stored in the following directory?

Do this import in a way that `Book` class name has no prefix. This doesn't mean that we can not have dots in our statment.   
After, create a new book, named `this_book`, with `name=prep-course`, `author=ldssa` and `year=2020`.   
__Hint__: Use `from _ import _ ` statment.

```python
!tree library
```

<!-- grade_id: cell-8dafeb32a322444f, points: 0, grade: false, solution: true -->
```python
#from _ import _
# YOUR CODE HERE
raise NotImplementedError()
#this_book = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9ac799d1b55a4cba, points: 0, grade: true, solution: false -->
```python
assert Book, "Book class was not imported with sucess"
assert isinstance(this_book, Book), "this_book is not of type Book"
assert this_book.name == "prep-course", "this_book has the wrong name"
assert this_book.author == "ldssa", "this_book has the wrong author"
assert this_book.year == 2020, "this_book has the wrong year"
```

### Q8 What is a package?

a) A folder;   
b) A set of modules that share the same directory;

<!-- grade_id: cell-17731565ade4c8d8, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ed761f5a351a7094, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q8, "Wrong answer."
```

### Q9 Should all modules that belong to a package be dedicated to the same functionality?

a) yes   
b) no

<!-- grade_id: cell-8df2491f6b0042a0, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-839cdbbdbfebce43, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q9, "Wrong answer."
```

### Q10 What is the mandatory file in a directory with modules to be considered a package?

a) `setup.py`   
b) `README.md`   
c) `LICENSE`   
d) `__init__.py`

<!-- grade_id: cell-471e5327dcc8e7b7, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9d78f1dc248bb974, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q10, "Wrong answer."
```

### Q11 What is the unix command used to manage packages in a python environment?

a) cat   
b) pip   
c) python   
d) tree   

<!-- grade_id: cell-6d4d775081c8be95, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-342a8848d6c5e6fe, points: 0, grade: true, solution: false -->
```python
assert utils.encode_answer(answer) == utils.q11, "Wrong answer."
```