# SLU03 - Exercise notebook

## Start by importing these packages

```python
# Just for evaluating the results.
import math
import json
import hashlib

#For extra visualization
import matplotlib.pyplot as plt
```

## Exercise 1: Conditionals and boolean algebra



## 1.1) Comparing Strings

In this exercise we have a series of strings that we want to sort from **smaller to larger**.

The answer should be a list where is **each element is a string that has lower value than the next string**:

For example: if we sort the strings "DDD", "XXX" and "AAA" by **ascending order of value** and insert them into a list the solution should be `["AAA", "DDD", "XXX"]`.


### 1.1.1) Sorting bananas

Sort the strings:

- `"banana"`
- `"ana"`
- `"ANA"`
- `"bananas"`
- `"BANANA"`

<!-- grade_id: cell-a5bd49d98abf2410, points: 0, grade: false, solution: true -->
```python
# Create a list with the strings above sorted in ascendant order of value.
# Assign the list to variable banana_list.

#Make sure that you type the strings exactly.

# banana_list = ...
# YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: cell-6b465d1d500666a0, points: 1, grade: true, solution: false -->
```python
banana_hash_0 = '482673bdd029354608338c077b71e987ba5571dfe7cd87329c50df73eae34fac'
banana_hash_1 = '82379da710fc913d545b2d3ea7c6b7a48e5cc9f3c8c7f63a7927be3153325109'
banana_hash_2 = '24d4b96f58da6d4a8512313bbd02a28ebf0ca95dec6e4c86ef78ce7f01e788ac'
banana_hash_3 = 'b493d48364afe44d11c0165cf470a4164d1e2609911ef998be868d46ade3de4e'
banana_hash_4 = 'e4ba5cbd251c98e6cd1c23f126a3b81d8d8328abc95387229850952b3ef9f904'

assert isinstance(banana_list, list), "Review the data type of banana_list."
assert len(banana_list) == 5, "The number of strings is incorrect."

assert banana_hash_0 == hashlib.sha256(bytes(banana_list[0], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert banana_hash_1 == hashlib.sha256(bytes(banana_list[1], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert banana_hash_2 == hashlib.sha256(bytes(banana_list[2], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert banana_hash_3 == hashlib.sha256(bytes(banana_list[3], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert banana_hash_4 == hashlib.sha256(bytes(banana_list[4], encoding='utf8')).hexdigest(), "At least one value is incorrect."
print("Your solution is correct!")
```

---

### 1.1.2) Sorting numbers as string

Sort the strings:

- `"1"`
- `"801"`
- `"02"`
- `"1000"`

<!-- grade_id: cell-cf0b9cd761180962, points: 0, grade: false, solution: true -->
```python
# Create a list with the strings above sorted in ascendent order of value.
# Assign the list to variable string_number_list.

#Make sure that you type the strings exactly.

# string_number_list = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ac16bba863ed00af, points: 1, grade: true, solution: false -->
```python
number_hash_0 = 'a953f09a1b6b6725b81956e9ad0b1eb49e3ad40004c04307ef8af6246a054116'
number_hash_1 = '6b86b273ff34fce19d6b804eff5a3f5747ada4eaa22f1d49c01e52ddb7875b4b'
number_hash_2 = '40510175845988f13f6162ed8526f0b09f73384467fa855e1e79b44a56562a58'
number_hash_3 = '096012b7ebcaf56d1d63b2784d2b2bbdeae080d72ad6bd1b9f7018e62a3c37d0'

assert isinstance(string_number_list, list), "Review the data type of string_number_list."
assert len(string_number_list) == 4, "The number of strings is incorrect."

assert number_hash_0 == hashlib.sha256(bytes(string_number_list[0], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert number_hash_1 == hashlib.sha256(bytes(string_number_list[1], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert number_hash_2 == hashlib.sha256(bytes(string_number_list[2], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert number_hash_3 == hashlib.sha256(bytes(string_number_list[3], encoding='utf8')).hexdigest(), "At least one value is incorrect."
```

---

## 1.2)  Order of precedence

In this exercise we have a boolean expression with a series of operators. Use the **table of precedence to identify the order in which the operations are executed**. Each operator is identified with a letter from `A` to `G`.
The solution should be a **list** where the first element is the letter of the first operation to be performed, the second element is the second operation to be performed, and so on until the last operation to be performed.

The result should resemble something like `["A", "B", "D", "G", ...]`

The string is:

```python
#You don't need to execute this cell!

variable = 3 != 1+1 and (2**3 > 1 or False)
#        A   B   C   D    E   F   G
```

The letters that identify the operations are:
- A: `=`
- B: `!=`
- C: `+`
- D: `and`
- E: `**`
- F: `>`
- G: `or`

<!-- grade_id: cell-ca00a15a43742c1c, points: 0, grade: false, solution: true -->
```python
# Create a list named operators_list that contains the identifiers of the operators sorted by the execution order.
# Make sure that you type the strings exactly.

# operators_list = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-43f5849625865cc0, points: 1, grade: true, solution: false -->
```python
operators_hash_0 = 'a9f51566bd6705f7ea6ad54bb9deb449f795582d6529a0e22207b8981233ec58'
operators_hash_1 = 'f67ab10ad4e4c53121b6a5fe4da9c10ddee905b978d3788d2723d7bfacbe28a9'
operators_hash_2 = '333e0a1e27815d0ceee55c473fe3dc93d56c63e3bee2b3b4aee8eed6d70191a3'
operators_hash_3 = '6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d'
operators_hash_4 = 'df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c'
operators_hash_5 = '3f39d5c348e5b79d06e842c114e6cc571583bbf44e4b0ebfda1a01ec05745d43'
operators_hash_6 = '559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd'

assert isinstance(operators_list, list), "Review the data type of operators_list."
assert len(operators_list) == 7, "The number of strings is incorrect."

assert operators_hash_0 == hashlib.sha256(bytes(operators_list[0], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert operators_hash_1 == hashlib.sha256(bytes(operators_list[1], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert operators_hash_2 == hashlib.sha256(bytes(operators_list[2], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert operators_hash_3 == hashlib.sha256(bytes(operators_list[3], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert operators_hash_4 == hashlib.sha256(bytes(operators_list[4], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert operators_hash_5 == hashlib.sha256(bytes(operators_list[5], encoding='utf8')).hexdigest(), "At least one value is incorrect."
assert operators_hash_6 == hashlib.sha256(bytes(operators_list[6], encoding='utf8')).hexdigest(), "At least one value is incorrect."
print("Your solution is correct!")
```

---

## 1.3)  `if-elif-else` statements

In this exercise we'll review `if-elif-else` statements. Each code will have inputs and `if-elif-else` statements. The goal is to **follow the flow** of the `if-elif-else` statements and **predict the value of certain variable afect execution**.

**Don't copy and execute the code**. If you can't follow these statements you won't be able to do the next exercises.

### 1.3.1) Reckless driving

Given the code:
```python
traffic_light = "Red"

if traffic_light == "Green":
    action = "Stop."
elif traffic_light == "Red":
    action = "Go forward."
else:
    action = "Turn right."

```

What is the value of `action`?

<!-- grade_id: cell-cce49b7314add616, points: 0, grade: false, solution: true -->
```python
# Assign to variable action the value it would have if the code above was executed. 

# action = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-97305f11f776d1a5, points: 1, grade: true, solution: false -->
```python
action_hash = 'aa4615f672c620de79597e6715674bdff165ea4764e73e9afa931c5bbb44439d'

assert isinstance(action, str), "Variable action should be a string."
assert action_hash == hashlib.sha256(bytes(action, encoding='utf8')).hexdigest(), "The value of variable action is incorrect."
print("Your solution is correct!")
```

---

### 1.3.2)  Cinema tickets

Given the code:
```python
age = 21
is_student = True
day_of_week = "Wednesday"
full_price = 10.00

if is_student and age < 18:
    ticket_price = full_price / 2
elif is_student and age >= 18:
    ticket_price = full_price * 3 / 4
elif day_of_week == "Monday":
    ticket_price = full_price * 2 / 3   
else:
    ticket_price = full_price
```

What is the value of `ticket_price`?

<!-- grade_id: cell-89716375f907d678, points: 0, grade: false, solution: true -->
```python
# Assign to variable ticket_price the value it would have if the code above was executed. 

# ticket_price = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1350367040ad9b58, points: 1, grade: true, solution: false -->
```python
assert isinstance(ticket_price, float), "Review the data type of ticket_price."
assert math.isclose(ticket_price, 7.5, abs_tol=0.005), "The value of ticket_price is incorrect."
print("Your solution is correct!")
```

---

### 1.3.3) Shopping for groceries

Given the code:
```python
groceries = {"eggs": 3, "milk": 4, "tomatoes": 4, "soaps": 1, "carrots": 6}
budget = 10

if len(groceries) > 4:
    if groceries["eggs"] > 6 and groceries["milk"] > 2:
        budget = 12
    elif groceries["tomatoes"] <= 2 or groceries["carrots"] < 3:
        budget = 32
    elif groceries["soaps"] < 4 or groceries["eggs"] > 6:
        budget = 40
    else:
        budget = 18
else:
    if groceries["eggs"] < 6 and groceries["milk"] <= 4:
        budget = 6
    elif groceries["tomatoes"] < 2 or groceries["carrots"] < 3:
        budget = 4
    else:
        budget = 8
```

What is the value of `budget`?

<!-- grade_id: cell-373d03ac81cc1129, points: 0, grade: false, solution: true -->
```python
# Assign to variable budget the value it would have if the code above was executed. 

# budget = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7b0857b8cabfbeab, points: 1, grade: true, solution: false -->
```python
budget_hash = 'd59eced1ded07f84c145592f65bdf854358e009c5cd705f5215bf18697fed103'

assert isinstance(budget, int), "Review the data type of budget."
assert budget_hash == hashlib.sha256(json.dumps(budget).encode()).hexdigest(), "The value of variable budget is incorrect."
print("Your solution is correct!")
```

---

# Exercise 2: `while` and `for` loops

In this exercise you can use `while` and/or `for` loops to perform the tasks. It's up to decide which loop(s) you need.

A tip: you won't need `if-elif-else` statements here.

## 2.1) Multiplying by 2

Given base 2, what is the **smallest exponent** that gives a **value larger than 13210**?

<!-- grade_id: cell-68e58a91a765c6cb, points: 0, grade: false, solution: true -->
```python
#Assign the exponent value on variable multi_counter.

# multi_counter = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b09b6a3df43346e0, points: 2, grade: true, solution: false -->
```python
multi_counter_hash = '8527a891e224136950ff32ca212b45bc93f69fbb801c3b1ebedac52775f99e61'

assert isinstance(multi_counter, int), "Review the data type of multi_counter."
assert multi_counter_hash == hashlib.sha256(json.dumps(multi_counter).encode()).hexdigest(), "The value of variable multi_counter is incorrect."
print("Your solution is correct!")
```

## 2.2) Summing  all positive integers until 100

Calculate the sum of all positive integers until 100 (`1 + 2 + 3 +...+ 100`).

<!-- grade_id: cell-0f95f2283fa7b5ed, points: 0, grade: false, solution: true -->
```python
# Assign the sum of all positive integer until 100 to gauss_sum.

#gauss_sum =
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b15473e30437404d, points: 2, grade: true, solution: false -->
```python
gauss_sum_hash = '3f95b1b8a32c2c0251dfdbc3c8a30aab6d6e680cf0ef03e8af84a65dff0c4a85'

assert isinstance(gauss_sum, int), "Review the data type of gauss_sum."
assert gauss_sum_hash == hashlib.sha256(json.dumps(gauss_sum).encode()).hexdigest(), "The value of variable gauss_sum is incorrect."
print("Your solution is correct!")
```

---

## 2.3) Popsicle in the freezer

In this exercise we'll be calculating the temperature of a popsicle in a freezer. Once the popsicle is in the freezer it starts to cooldown until it has the same temperature as the freezer. But this decrease in temperature is not linear over time. It depends on the difference of temperature between the popsicle and the freezer.

The initial temperature of the popsicle is stored in `temp_popsicle[0]`. The **next** temperature of the popsicle is given by

`temp_popsicle[1] = (temp_freezer - temp_popsicle[0]]) * dt + temp_popsicle[0]`

where:

- `temp_popsicle[1]` is the next temperature of the popsicle (after `dt` passed)
- `temp_popsicle[0]` is the current temperature of the popsicle
- `temp_freezer` is the temperature of the freezer. It is always constant.
- `dt` is the time between observations.


Your task is to use the values available and calculate the next temperature and **append it** to `temp_popsicle`. It becomes a list of size 2. The first element is the original temperature and the second element is the new calculated temperature. Then repeat the process by adding 1 on the indices of the equation so that for the second iteration `temp_popsicle[1]` becomes `temp_popsicle[2]` and `temp_popsicle[0]` becomes `temp_popsicle[1]`. Do this until you get a list with **length 80**.

The value of temperature on index `i` of `temp_popsicle` is given by: 

`temp_popsicle[i] = (temp_freezer - temp_popsicle[i-1]]) * dt + temp_popsicle[i-1]`

**Important**: everytime you want to execute your solution you need to execute the **next cell before**. If you don't, the `temp_popsicle` list will keep getting larger and larger size and you won't have the right result.

The result should evolve like this:

`temp_popsicle = [20]`

becomes

`temp_popsicle = [20 17.0]`
 
and the equation becomes `temp_popsicle[2] = (temp_freezer - temp_popsicle[1]]) * dt + temp_popsicle[1]` resulting in:
 
`temp_popsicle = [20 17.0 14.3]`
 
Then you update the equation and calculate the next temperature and so on until you have a list of **length 80**.

```python
#Parameters of the exercise. Anytime you execute the exercise, execute this first.
temp_popsicle = [20.0]
temp_freezer = -10
dt = 0.1
```

<!-- grade_id: cell-4a8b61ca7fd9597b, points: 0, grade: false, solution: true -->
```python
#Append the new temperature to temp_popsicle. Use the new temperature to calculate the next temperature and so on.

# temp_popsicle= [...]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0c8fd7edbbfaf642, points: 2, grade: true, solution: false -->
```python
#Everytime you want to execute your code, execute the one above first!!!

popsicle_hash = '48449a14a4ff7d79bb7a1b6f3d488eba397c36ef25634c111b49baf362511afc'

assert isinstance(temp_popsicle, list), "Review the data type of temp_popsicle."
assert popsicle_hash == hashlib.sha256(json.dumps(len(temp_popsicle)).encode()).hexdigest(), "The length of temp_popsicle is incorrect. Make sure to execute the parameter cell first!"
assert math.isclose(temp_popsicle[0], 20.0, abs_tol=0.1), "The first value of temp_popsicle is incorrect."
assert math.isclose(temp_popsicle[1], 17.0, abs_tol=0.1), "The first temperature update is incorrect."
assert math.isclose(temp_popsicle[2], 14.3, abs_tol=0.1), "The second temperature update is incorrect."
assert math.isclose(temp_popsicle[-1], -9.9, abs_tol=0.4), "The final temperature is incorrect."
print("Your solution is correct!")
```

---

### Optional visualization (Not scored)

Below a simple code that creates a plot of the temperature of the popsicle over time. This is only a sneak peak of the visualizations that are possible with matplotlib.

If you are having problems with code below, you can ignored it. 

```python
# I encourage you to, after completing the exercises, play with this code:
# Change the colors and labels, use others lists as input, break the code.


%matplotlib inline
plt.rcParams["figure.figsize"] = (10,7)
ax=plt.gca()

x = range(len(temp_popsicle))
y = temp_popsicle

plt.plot(x,y, '-', color='tab:blue')

# General settings
plt.xlim(x[0], x[-1])
plt.ylim(-11, 20)
plt.ylabel("Temperature of popsicle (Cº)",fontsize='xx-large')
plt.xlabel('Time (min)',fontsize='xx-large')


plt.grid(True, 'major', 'y', ls='--', lw=.5, c='k', alpha=.3)
# Show graphic
plt.show()
```

---

# 3) Using conditionals and loops together

In this exercise, you'll need to use `if-elif-else` but also loops to solve the tasks.

<img src="./media/fun_begins.gif" />



## 3.1) The Collatz conjecture

From [Wikipedia](https://en.wikipedia.org/wiki/Collatz_conjecture):

> The Collatz conjecture is a conjecture in mathematics that concerns a sequence defined as follows: start with any positive integer n. Then each term is obtained from the previous term as follows: if the previous term is even, the next term is one half of the previous term. If the previous term is odd, the next term is 3 times the previous term plus 1. The conjecture is that no matter what value of n, the sequence will always reach 1. 

We'll put the conjecture to the test. We'll start the task with an initial value of `n = 134`.

Next we'll apply an operation over `n` depending on its value:

If `n` is **even** the next value of `n` is given by `n / 2`.

If `n` is **odd** the next value of `n` is given by `3*n + 1`.

We'll repeat this process until `n` is equal to 1.

We want to know **how many operations are performed until `n` reaches 1** and assign this value to variable `collatz_operations`.

<!-- grade_id: cell-ca40d76b6b0fd084, points: 0, grade: false, solution: true -->
```python
#Assign the number of operations required to reach 1 to variable collatz_operations.

# collatz_operations =
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-822b5ec7dfe02776, points: 2, grade: true, solution: false -->
```python
collatz_operations_hash = '59e19706d51d39f66711c2653cd7eb1291c94d9b55eb14bda74ce4dc636d015a'

assert isinstance(collatz_operations, int), "The number of operations should be an integer."
assert collatz_operations_hash == hashlib.sha256(json.dumps(collatz_operations).encode()).hexdigest(), "The number of operations is incorrect."
print("Your solution is correct!")
```

---

## 2.2) Football championship points

In a particular football league a certain number of matches were played. A team wins if it has more goals than their opponents and the teams tie if the number of goals are equal. The team that wins the match gets 3 points. The losing team gets 0 points. In case of a tie, both teams get 1 point.

The result of the matches is stored in `matches_outcome`, where each element is the match output in terms of goals.
The goals of each team are placed after the name of the team.

For example, `("Deadly Reptiles",1,"Short Clowns",3)` means that the Deadly Reptiles scored 1 goal and the Short Clowns scored 3 goals.

Update the dictionary `club_points` with the final number of points of each team.

**Important**: everytime you want to execute your solution you need to execute the next cell before. If you don't, teams will keep getting more and more points and you won't have the right result.

```python
matches_outcome = (("Deadly Reptiles",1,"Short Clowns",3),
                   ("Breakfast Fallout",1,"Donkey Attackers",1),
                   ("Deadly Reptiles",2,"Breakfast Fallout",0),
                   ("Short Clowns",1,"Donkey Attackers",1),
                   ("Short Clowns",0,"Breakfast Fallout",0),
                   ("Donkey Attackers",4,"Deadly Reptiles",3)
                  )

club_points = {"Deadly Reptiles":0, "Short Clowns":0, "Breakfast Fallout":0, "Donkey Attackers":0}
```

<!-- grade_id: cell-0a458560de5f3c4e, points: 0, grade: false, solution: true -->
```python
#Iterate over every match and attribute the points to the respective teams.
#Everytime you want to execute this code, execute the one above first!!!

# club_points = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-65f0dc204b4b0812, points: 2, grade: true, solution: false -->
```python
#Everytime you want to execute your code, execute the one above first!!!

reptile_hash = '4e07408562bedb8b60ce05c1decfe3ad16b72230967de01f640b7e4729b49fce'
clowns_hash = 'ef2d127de37b942baad06145e54b0c619a1f22327b2ebbcfbec78f5564afe39d'
fallout_hash = 'd4735e3a265e16eee03f59718b9b5d03019c07d8b6c51f90da3a666eec13ab35'
attackers_hash = 'ef2d127de37b942baad06145e54b0c619a1f22327b2ebbcfbec78f5564afe39d'

assert isinstance(club_points, dict), "Variable club_points should be a dictionary."
assert reptile_hash == hashlib.sha256(json.dumps(club_points["Deadly Reptiles"]).encode()).hexdigest(), "The number of points for the Deadly Reptiles is incorrect."
assert clowns_hash == hashlib.sha256(json.dumps(club_points["Short Clowns"]).encode()).hexdigest(), "The number of points for the Deadly Reptiles is incorrect."
assert fallout_hash == hashlib.sha256(json.dumps(club_points["Breakfast Fallout"]).encode()).hexdigest(), "The number of points for the Deadly Reptiles is incorrect."
assert attackers_hash == hashlib.sha256(json.dumps(club_points["Donkey Attackers"]).encode()).hexdigest(), "The number of points for the Deadly Reptiles is incorrect."
print("Your solution is correct!")
```

---

# 3) List comprehension

In this exercise we'll perform some tasks that are simpler to do with list comprehension. 

## 3.1) Filter and grow

Given the list `["Lighting", 23, 1.3, "potatoes", [1, 2, 3], 4.5, 7]` we want to create a list with **only the numeric values and each value should be multiplied by 3.**
You should solve this with one line of code.

```python
original_list = ["Lighting", 24, 1.21, "potatoes", [1, 2, 3], 4.5, 7]
```

<!-- grade_id: cell-d2b9abfbcaf35a41, points: 0, grade: false, solution: true -->
```python
# Create a list comprehension that iterates over original_list.
# The result should be a list assigned to processed_list.
# Execute the cell above 

# processed_list = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9f9722689d8a954e, points: 2, grade: true, solution: false -->
```python
processed_hash = '4b227777d4dd1fc61c6f884f48641d02b4d121d3fd328cb08b5531fcacdabf8a'

assert isinstance(processed_list, list), "Review the data type of temp_popsicle."
assert processed_hash == hashlib.sha256(json.dumps(len(processed_list)).encode()).hexdigest(), "The length of processed_list is incorrect."
assert math.isclose(processed_list[0], 72, abs_tol=0.1), "The some values of processed_list are incorrect."
assert math.isclose(processed_list[2], 13.5, abs_tol=0.1), "The some values of processed_list are incorrect."
print("Your solution is correct!")
```

---

## 3.2) Symmetric matrix

Create the following list of lists (matrix):

$
\begin{bmatrix}
\begin{bmatrix}0 & 1 & 2 & 3 & 4\end{bmatrix},\\
\begin{bmatrix}1 & 2 & 3 & 4 & 5\end{bmatrix},\\
\begin{bmatrix}2 & 3 & 4 & 5 & 6\end{bmatrix},\\
\begin{bmatrix}3 & 4 & 5 & 6 & 7\end{bmatrix},\\
\begin{bmatrix}4 & 5 & 6 & 7 & 8\end{bmatrix}\space\\
\end{bmatrix}
$

<!-- grade_id: cell-6cc7446a5944ed73, points: 0, grade: false, solution: true -->
```python
# Create the list of lists above using list comprehension.
# The result should be a list assigned to matrix_list. 

# matrix_list = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-40aece1a4ce01dcb, points: 2, grade: true, solution: false -->
```python
matrix_hash = 'ef2d127de37b942baad06145e54b0c619a1f22327b2ebbcfbec78f5564afe39d'

assert isinstance(matrix_list, list), "Review the data type of temp_popsicle."
assert matrix_hash == hashlib.sha256(json.dumps(len(matrix_list)).encode()).hexdigest(), "The length of matrix_list is incorrect."
assert matrix_hash == hashlib.sha256(json.dumps(len(matrix_list[0])).encode()).hexdigest(), "The length of the first element of matrix_list is incorrect."
assert matrix_hash == hashlib.sha256(json.dumps(matrix_list[2][3]).encode()).hexdigest(), "Some values are incorrect."
assert matrix_hash == hashlib.sha256(json.dumps(matrix_list[3][2]).encode()).hexdigest(), "Some values are incorrect."
print("Your solution is correct!")
```

---

### Extra fact:

There is a [story](http://mathandmultimedia.com/2010/09/15/sum-first-n-positive-integers/) about how [Gauss](https://en.wikipedia.org/wiki/Carl_Friedrich_Gauss#Anecdotes) discovered a formula to calculate the sum of the first `n` positive integers. If there is a formula or a simple calculation that you can perform to avoid extensive code, why not using it instead? 