# SLU12 - Git Basics - Exercise notebook

Start by importing the following

```python
# Just for evaluating the results.
import hashlib
```

# Exercise 1: Version Control

What is one of the most important reasons why version control software are used in the industry?
- Option A: It is easier than sharing code by e-mail
- Option B: It allows several people to work on the same codebase in a concurrent way and allows them to resolve conflicts if they occur
- Option C: It allows people to create public repositories and share their code with everyone

<!-- grade_id: cell-71ed0e811c050c93, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the option value of the option that you think is correct to the variable exercise_1.
# For example, if you think that option A is correct, do: exercise_1 = "A"

### BEGIN SOLUTION
exercise_1 = "B"
### END SOLUTION
# exercise_1 =
```

<!-- grade_id: cell-ac103768b1b222f7, points: 0, grade: true, solution: false -->
```python
exercise_1_hash = 'df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c'
assert isinstance(exercise_1,str), "The variable should be a string."
assert exercise_1_hash == hashlib.sha256(bytes(exercise_1, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

# Exercise 2: Staging Area

What is the main advantage of using the staging area?

- Option A: It allows us to make atomic and logical commits
- Option B: It allows us to separate the workspace from the local repository
- Option C: It has a copy of the remote repository

<!-- grade_id: cell-ec8605dfc3b02035, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the option value of the option that you think is correct to the variable exercise_2.

### BEGIN SOLUTION
exercise_2 = "A"
### END SOLUTION
# exercise_2 =
```

<!-- grade_id: cell-6ee0d9ea58491a71, points: 0, grade: true, solution: false -->
```python
exercise_2_hash = '559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd'
assert isinstance(exercise_2,str), "The variable should be a string."
assert exercise_2_hash == hashlib.sha256(bytes(exercise_2, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

# Exercise 3: Git Add

What is the correct command to use when you want to stage all the changes in the workspace?

- Option A: `git add --all`
- Option B: `git add .`
- Option C: `git add workspace`

<!-- grade_id: cell-b6d48cc3f1b95e0d, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the option value of the option that you think is correct to the variable exercise_3.

### BEGIN SOLUTION
exercise_3 = "B"
### END SOLUTION
# exercise_3 =
```

<!-- grade_id: cell-15ef0dc5325e7dd6, points: 0, grade: true, solution: false -->
```python
exercise_3_hash = 'df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c'
assert isinstance(exercise_3,str), "The variable should be a string."
assert exercise_3_hash == hashlib.sha256(bytes(exercise_3, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

# Exercise 4: Git Status

You have 10 files in a repository. Of those 10, you've already committed in the past 6 of them. Of those 6, you have now staged 2 files. How many files are not being tracked and how many files are not staged?

- Option A: 10 and 6
- Option B: 4 and 6
- Option C: 4 and 4

<!-- grade_id: cell-5c3ec2bc3145164e, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the option value of the option that you think is correct to the variable exercise_4.

### BEGIN SOLUTION
exercise_4 = "C"
### END SOLUTION
# exercise_4 =
```

<!-- grade_id: cell-82affb167d780260, points: 0, grade: true, solution: false -->
```python
exercise_4_hash = '6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d'
assert isinstance(exercise_4,str), "The variable should be a string."
assert exercise_4_hash == hashlib.sha256(bytes(exercise_4, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

# Exercise 5: Git Diff

After you run git diff on a file, the chuck header is @@ -1,4 +1,7 @@. Given that the number of lines removed is 2, what can you say about the number of lines added?

- Option A: It is a number between 5 and 7
- Option B: It is a number between 1 and 7
- Option C: It is a number between 1 and 5 

<!-- grade_id: cell-8a385c025ec6dafa, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the option value of the option that you think is correct to the variable exercise_5.

### BEGIN SOLUTION
exercise_5 = "C"
### END SOLUTION
# exercise_5 =
```

<!-- grade_id: cell-e400aca09757cc0f, points: 0, grade: true, solution: false -->
```python
exercise_5_hash = '6b23c0d5f35d1b11f9b683f0b0a617355deb11277d91ae091d399c655b87940d'
assert isinstance(exercise_5,str), "The variable should be a string."
assert exercise_5_hash == hashlib.sha256(bytes(exercise_5, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

# Exercise 6: Git Log

You want to see the last 5 commits on your repo. What command should you run?

- Option A: `git log --last 5`
- Option B: `git log -n 5`
- Option C: `git log 5`

<!-- grade_id: cell-989115f1ea143a4d, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the option value of the option that you think is correct to the variable exercise_6.

### BEGIN SOLUTION
exercise_6 = "B"
### END SOLUTION
# exercise_6 =
```

<!-- grade_id: cell-094ec2cd7745cd9b, points: 0, grade: true, solution: false -->
```python
exercise_6_hash = 'df7e70e5021544f4834bbee64a9e3789febc4be81470df629cad6ddb03320a5c'
assert isinstance(exercise_6,str), "The variable should be a string."
assert exercise_6_hash == hashlib.sha256(bytes(exercise_6, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

# Exercise 7: .gitignore

You do not want to commit any .md file except for the README. What two lines should you include in your .gitignore?

- Option A: *.md and !README.md
- Option B: .md and !README.md
- Option C: *.md and !README

<!-- grade_id: cell-dde4b8e88ebb208b, points: 0, grade: false, solution: true -->
```python
# The task is:
# Assign the option value of the option that you think is correct to the variable exercise_7.

### BEGIN SOLUTION
exercise_7 = "A"
### END SOLUTION
# exercise_7 =
```

<!-- grade_id: cell-b4030ab1152948cb, points: 0, grade: true, solution: false -->
```python
exercise_7_hash = '559aead08264d5795d3909718cdd05abd49572e84fe55590eef31a88a08fdffd'
assert isinstance(exercise_7,str), "The variable should be a string."
assert exercise_7_hash == hashlib.sha256(bytes(exercise_7, encoding='utf8')).hexdigest(), "The value of the option is incorrect."
```

# Exercise 8: Git Workflow

This is the last exercise of the notebook. It is basically a step of instructions for you to do and see for yourself how some git commands work:
- Save this file
- On the command line, run this command: `cd ~/projects/ds-prep-workspace`
- Run the git status command
- Check the output, you should see that the Exercise Notebook file is Modified
- Run the git add command to add just the Exercise Notebook file
- Run the git commit command with the option that allows you to write the commit message in the command. Write the following message: "This is the first commit that I truly understand! Hopefully the first of many!"

The end! Feel free to say this on the week 7 channel: "@Miguel D - Instructor, Setup I challenge you for a fight! That sword is no match for me!".