# SLU06 - File & String handling

Now we're going to test how well you understood the learning notebook.

Also, this notebook is going to often require some googling skills. It's very important to learn to google anything you don't remember or don't know how to do.

A small hint: list comprehensions might make it easier to solve sove of the exercises

### Exercise 1
- Open a local file called 'assignment.txt'. Store it in a variable called 'f'
- Read the whole file and store it in a variable called 'text'. The variable should consist of only one string representing all the lines from the file

<!-- grade_id: cell-5c616f01a6b14b63, points: 0, grade: false, solution: true -->
```python
# f = ...
# text = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e2d149b3016dc2fc, points: 1, grade: true, solution: false -->
```python
assert isinstance(text, str), "Are you sure you read the whole file as one string (not as list of strings)?"
assert len(text) == 727, "The length of the string doesn't match"
assert text[0]=='O', "Did you read the correct file?"
```

### Exercise 2
- Move the read cursor to the beginning of the file
- Read each line of the file and store it in a variable called lines. This line should be a list of strings.

<!-- grade_id: cell-78e73bd6fb0635fc, points: 0, grade: false, solution: true -->
```python
# ...
# lines = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d1ac86663d79cef7, points: 1, grade: true, solution: false -->
```python
assert len(lines) != 0, "Did you move the cursor to the beginning of the file?"
assert isinstance(lines, list), "Are you sure you stored the file as a list of strings?"
assert len(lines) == 13, "The number of string elements in the list doesn't match"
assert lines[0]=='On an exceptionally hot evening early in July a young man came out of\n', "Are you sure you read the correct file?"
```

### Exercise 3
Let's preprocess our file.
- For each string in the list, remove all the newline characters and save the result in the same variable

<!-- grade_id: cell-d3af9181b11c5b5b, points: 0, grade: false, solution: true -->
```python
# lines = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-62eba51f14c12c49, points: 1, grade: true, solution: false -->
```python
assert isinstance(lines, list), "Lines has to be a list of strings"
assert len(lines) == 13, "The number of strings in the list doesn't match"
assert len(lines[0])==69, "Are you sure you removed the newline characters?"
```

### Exercise 4
Remove all the empty lines:

<!-- grade_id: cell-cca50a46d15fd13d, points: 0, grade: false, solution: true -->
```python
# lines = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-048486c832b1aeb9, points: 2, grade: true, solution: false -->
```python
assert len(lines) == 11, "Are you sure that you removed all the empty strings?"
```

### Exercise 5
Concatenate all the strings in the list into one long string separating them with a space symbol

<!-- grade_id: cell-121fbc38e772e04a, points: 0, grade: false, solution: true -->
```python
# lines = ...
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d6a132cb0e86cc5d, points: 2, grade: true, solution: false -->
```python
assert isinstance(lines, str), "It's not a string"
assert len(lines) == 725, "The length of the string doesn't match"
```

### Exercise 6
Convert the string to a list of words and save it in a variable called 'words'.
Hint: Use space symbol as a separator

<!-- grade_id: cell-0023979f541f52e0, points: 0, grade: false, solution: true -->
```python
# words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-84bffa431d75f2f6, points: 2, grade: true, solution: false -->
```python
assert isinstance(words, list), "It's not a list"
assert isinstance(words[0], str), "The list has to consist of strings"
assert len(words) == 132, "The number of words doesn't match"
```

### Exercise 7
Convert all the words to lowercase

<!-- grade_id: cell-cb00efd631ee3d37, points: 0, grade: false, solution: true -->
```python
# words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ab311a01acff014c, points: 1, grade: true, solution: false -->
```python
assert all([word.islower() for word in words]), "Not all the words are lowercase"
```

### Exercise 8
In Natural Language Processing we usually want to remove some words that have no meaning. They are usually called 'stop words'.

We're not dealing with NLP yet, but let's also remove such words.
- remove the words that are in the following list:
['on','an','in','a','of','the','and']

<!-- grade_id: cell-064d872c33ebafdd, points: 0, grade: false, solution: true -->
```python
# stop_words = ['on','an','in','a','of','the','and']
# words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e51d4948d757891, points: 2, grade: true, solution: false -->
```python
assert len(words) == 100, "The number of words doesn't match"
```

### Exercise 9
Keep only unique words in this list words created on exercise 8.

Store them in the same variable. It has to be a list of unique words

P.S., Don't remove any symbols in the strings! Here we assume that if the original string had both 'hello.' and 'hello', they are both unique and we keep both of them. 

Hint: You'll need some Googling on this question.

<!-- grade_id: cell-aa8cd00bd90721fb, points: 0, grade: false, solution: true -->
```python
# words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e354a40c7e837c45, points: 3, grade: true, solution: false -->
```python
assert isinstance(words, list), "We need to store the unique words in a list"
assert len(words) == 77, "The number of words doesn't match"
```

### Exercise 10
Create a function to find if a word has more characters than a given number:
- Create the function is_long
- Receive 2 parameters: a string and an integer
- The string represents the word that we're checking
- The integer represents the minimum length that word should have
- return True if length of the word >= integer. Return False otherwise

<!-- grade_id: cell-2cfca9c135f57b08, points: 0, grade: false, solution: true -->
```python
# def is_long(word, num):
#    ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f83d24ebe38a8c0e, points: 2, grade: true, solution: false -->
```python
assert is_long('Exercise', 5), "The function didn't pass the test for the following parameters: word='exercise', num=5"
assert not is_long('Exercise', 10), "The function didn't pass the test for the following parameters: word='exercise', num=10"
```

### Exercise 11
Apply is_long function to each word in 'words' variable you have created in Exercise 9.
Assume the integer parameter num = 5.
Store the result in a variable called long_words

<!-- grade_id: cell-d6886505b94e4ce5, points: 0, grade: false, solution: true -->
```python
# long_words = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8c31b9ac6bc23548, points: 1, grade: true, solution: false -->
```python
assert len(long_words) == 47, "The number of words doesn't match"
```

### Exercise 12
Create a function that calculates how many words in a list start with a given letter.
- Create the function start_with
- Accept 2 parameters: a list of words, a letter to check
- return the number of cases when this letter is the first letter in a word

<!-- grade_id: cell-2088771824f8e2a1, points: 0, grade: false, solution: true -->
```python
# def start_with(words_list, letter):
#    ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8bb4a1579925fe13, points: 2, grade: true, solution: false -->
```python
assert start_with(long_words, 'b') == 2, "The function didn't work for long_words and letter='b'"
assert start_with(words, 'a') == 5, "The function didn't work for words and letter='a'"
assert start_with(words, 'd') == 3, "The function didn't work for words and letter='d'"
```

### Exercise 13

This task is going to be a bit more difficult, but also pretty realistic. We're going to create our own encoder.    

Encoder is a program that transforms one format of data into another format. You're going to use them all the time on the academy. Now we're going to write a simple encoder by ourselves.

Our encoder is going to check if an sentence has some predifined words. Imagine that we have a lot of sentences and we want to understand if they are positive or negative.

Our idea is the following:

If the sentence has words like 'good', 'awesome', 'fantastic', 'hilarious' etc., it's positive. And it's negative otherwise.
Let's call the words like 'good' our vocabulary.

Create a function, which:
- is called encoder()
- gets 2 inputs: 1)sentence (string) 2)vocabulary (list of words)
- for each word in the vocabulary, checks whether this word is in the sentence.
- returns a list of length of vocabulary, but instead of words it has only zeros and ones. 1 means that this word was in the sentence, 0 otherwise.
Example:
`vocabulary = ['good', 'better', 'awesome', 'fantastic']
sentence = 'This day was fantastic. Tomorrow will be even better!'
encoder(sentence, vocabulary)`

Output:
> [0,1,0,1]

Explanation: 

'good' was not in the sentence => 0
'better' was in the sentence => 1
'awesome' was not in the sentence => 0
'fantastic' was in the sentence => 1'

Don't forget to lowercase the words in the sentence and in the vocabulary. Assume that all the words are separated by space symbol.
If you're stuck with this task, try to divide it in smaller parts. 

**A hint:**

we can divide this task into a few smaller steps.

1) We might firstly need to lowercase each word in the vocabulary

2) Split 'sentence' into a list of words. Let's call it "words_list"

3) Lowercase each word in this list.

4) For each word in our vocabulary, ask the program whether this word is in words_list. If yes, the answer is 1. If it's not, the answer is 0.

5) return the list of answers

<!-- grade_id: cell-24946c705e259797, points: 0, grade: false, solution: true -->
```python
# def encoder(sentence, vocabulary):
#    ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-28362929ccc11e01, points: 5, grade: true, solution: false -->
```python
assert isinstance(encoder('The day was good', ['good']), list), 'The function has to return a list'
assert len(encoder('The day was good', ['good'])) != 0, 'Your function returns an empty list'
assert isinstance(encoder('The day was good', ['good'])[0], int), 'The elements of the output list have to be integers'
assert len(encoder('The day was good', ['good', 'bad', 'awesome'])) == 3, 'Number of elements in the output list does not match the length of the vocabulary'
assert encoder('The day was good', ['good', 'bad', 'awesome']) == [1,0,0], "The function failed on the test: encoder('The day was good', ['good', 'bad', 'awesome'])"
```

### Exercise 14

`def function(a):
    assert isinstance(a, int)
    print(a + 1)`
    
P.S, isinstance function checks whether a variable has a specific format. In our case it checks whether variable a is in integer
    
What is going to happen if you call:

`function(3.5)`

Options:
- a. 4.5 will be printed
- b. AssertionError will be raised
- c. AssertionError will be printed
- d. Nothing. The code is not valid

Write the letter with the correct answer to a variable called 'answer' as a string.

For example,

`answer = 'a'`

Note: try to think before writing anything. If you don't know the right answer, check the learning material.

But please try not to guess the answer by testing the function itself (don't call it).

You're learning it for yourself, not for correct answers or some grades 🙂 

<!-- grade_id: cell-213d3992590c59fb, points: 0, grade: false, solution: true -->
```python
# answer = ''

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6ea081f9ee1904d5, points: 1, grade: true, solution: false -->
```python
assert isinstance(answer, str), 'The answer should be a string'
assert answer == 'b', 'The answer is wrong'
```

### Exercise 15

What's the output of the following line?

`'I am' + 17 + 'years old'`

Options:
- a. Name error
- b. 'I am 17 years old'
- c. Type error
- d. 'I am years old 17'

Save the answer in a variable called 'answer'. Example:

`answer = 'a'`

<!-- grade_id: cell-dcc00f6894cf84b2, points: 0, grade: false, solution: true -->
```python
# answer = ''

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9e2b046778a27c04, points: 1, grade: true, solution: false -->
```python
assert isinstance(answer, str), 'The answer should be a string'
assert answer == 'c', 'The answer is wrong'
```

### Exercise 16

What's the result of the following lines?

`text = ['I', 'am', 'happy', 'today']
print(text[4])`

Options:
- a. 'today'
- b. 'h'
- c. Type error
- d. Index error

Save the answer in a variable called 'answer'. Example:

`answer = 'a'`

<!-- grade_id: cell-c26141d1871dd231, points: 0, grade: false, solution: true -->
```python
# answer = ''

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-279c5d1dbdcbf4f2, points: 1, grade: true, solution: false -->
```python
assert isinstance(answer, str), 'The answer should be a string'
assert answer == 'd', 'The answer is wrong'
```

### Exercise 17

Create a function called join_words that accepts a list of strings called `list_strings` and a string `separator`.    
Join all strings on `list_strings` with string `separator` between all words to a new variable called `sentence`.    Return `sentence`.


If the operation goes wrong return the following string `"function inputs don't match the requirements"`

<!-- grade_id: cell-ef39f20f9c0de7f0, points: 0, grade: false, solution: true -->
```python
#def join_words(list_strings, separator):

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-726e1e35c284e034, points: 1, grade: true, solution: false -->
```python
assert isinstance(join_words(["this", "is", "a", "string"], " "), str), "String output is expected"
assert join_words(["this", "is", "a", "string"], " ") == "this is a string", "The answer is wrong for input=['This', 'is', 'a', 'string']"
assert join_words(["this", "is", False, "a", "string"], " ") == "function inputs don't match the requirements", "Error message was not returned"
```

### Exercise 18

Create a function called multiply_not_five, which:
- has an integer num as an input
- return the number multipled by 2
- if the number is equal to 5, raises an AssertionError

<!-- grade_id: cell-e778b52bc04648f8, points: 0, grade: false, solution: true -->
```python
#def multiply_not_five(num):

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-845a471dd5f6e303, points: 1, grade: true, solution: false -->
```python
assert multiply_not_five(10) == 20, 'Wrong answer for number = 10'
assert multiply_not_five(-5) == -10, 'Wrong answe for number = -5'
def check_answer(function):
    try:
        function(5)
        raise Exception("The function shouldn't work for number = 5")
    except AssertionError:
        pass
    except:
        raise Exception("You need to return an AssertionError, not another exception type")
check_answer(multiply_not_five)
```

### Exercise 19

The last exercise is pretty similar to exercise 18. 

Create a function called multiply_string, which:
- has a string as an input
- return the string repeated 10 times.E.g, multiply_string('a') = 'aaaaaaaaaa'
- **hint**: you might want to google how to repeat strings in python
- if the input is not a string, raise an Exception with some text output. For example, Exception("It's not a string, dude!")
- **hint**: you might want to google how to check whether a variable is a string in python

<!-- grade_id: cell-6f4c85c9880ef5d8, points: 0, grade: false, solution: true -->
```python
#def multiply_string(string):

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a055f00629e597a8, points: 1, grade: true, solution: false -->
```python
assert multiply_string('a') == 'aaaaaaaaaa', "The function didn't work for input='a'"
assert multiply_string('z') == 'zzzzzzzzzz', "The function didn't work for input='z'"
def check_answer(function):
    try:
        function(5)
        raise Exception("The function shouldn't work for integers")
    except:
        pass
check_answer(multiply_string)
```