# SLU02 - Data Structures

### Start by importing the following packages

```python
#used for evaluation
import hashlib
import json
import random
```

In this notebook the following is tested:
- Tuples   
- Lists   
- Dictionaries

## Exercise 1: Tuples  <a name="1"></a>

This exercise covers topics learned regarding tuples.

### 1.1) Create a tuple 

Create a tuple of __floats__ named `this_tuple` with size 5.

<!-- grade_id: cell-3a4c0b59df7bcf57, points: 0, grade: false, solution: true -->
```python
# this_tuple = ...

### BEGIN SOLUTION
this_tuple = (1., 2., 3., 4., 5.)
print(type(this_tuple))
print(len(this_tuple))
print(type(this_tuple[3]))
### END SOLUTION
```

Output:
```
<class 'tuple'>
5
<class 'float'>
```

<!-- grade_id: cell-b63230a58719fa25, points: 1, grade: true, solution: false -->
```python
assert isinstance(this_tuple, tuple), "Are you sure this_tuple is a tuple?"
assert len(this_tuple) == 5, "The length is not quite right."
assert isinstance(this_tuple[3], float), "Did you write floats?"
```

### 1.2) Index a tuple

Considering the following tuple:    

```python
color = ("red", "blue", "green", "yellow", "black", "white")
```

Using __negative__ indexing, assign the index of the element `"green"` to the variable `green_index`.

<!-- grade_id: cell-268b39f922ddc862, points: 0, grade: false, solution: true -->
```python
color = ("red", "blue", "green", "yellow", "black", "white")
#hint: the solution should be a number
#green_index = ...

### BEGIN SOLUTION
green_index = -4
print(color[green_index])
print(
    'green_index hashed:',
    hashlib.sha256(
        json.dumps(green_index).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
green
green_index hashed: e5e0093f285a4fb94c3fcc2ad7fd04edd10d429ccda87a9aa5e4718efadf182e
```

<!-- grade_id: cell-d3eb2f19b4f68632, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(green_index).encode()).hexdigest() == 'e5e0093f285a4fb94c3fcc2ad7fd04edd10d429ccda87a9aa5e4718efadf182e', "The index is not correct. Are you using negative indexing?"
```

### 1.3) Slice a tuple

Extract `("green", "blue", "red")` from tuple `color`.   
Assign the results to a variable called `rgb`. 

<!-- grade_id: cell-205bd03d74d9e9f5, points: 0, grade: false, solution: true -->
```python
color = ("red", "blue", "green", "yellow", "black", "white")
#Hint: use backwards slicing
#rgb = ...

### BEGIN SOLUTION
rgb = color[-4:-7:-1]
print(rgb)
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(rgb).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
('green', 'blue', 'red')
rgb hashed: 43da8949efc00c51f7d96130a25a6b902f6cfd6157817015ca6ee524e3085374
```

<!-- grade_id: cell-f17f899e57df23ad, points: 1, grade: true, solution: false -->
```python
assert isinstance(rgb, tuple), "Is your result a tuple?"
assert len(rgb) == 3, "You aren't selecting the correct number of elements."
assert rgb[0] == "green", "Check which elements you are selecting and their order."
assert hashlib.sha256(json.dumps(rgb).encode()).hexdigest() == '43da8949efc00c51f7d96130a25a6b902f6cfd6157817015ca6ee524e3085374'
```

###  1.4) Index a tuple of tuples

Considering the following tuple of tuples:

`random_numbers = ((1, 2, 3),(4, 5, 6),(7, 8, 9),(10, 11, 12))`

What is the right way to extract the number 8 from `random_numbers`?

a) `random_numbers[3][2]`   
b) `random_numbers[-1][-1]`   
c) `random_numbers[-2][1]`  
d) `random_numbers[-1][1]`   

<!-- grade_id: cell-7e6632972efca3db, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
### BEGIN SOLUTION
random_numbers = ( (1, 2, 3), (4, 5, 6), (7, 8, 9), (10, 11, 12) )
answer = "c"
print(answer)
print(random_numbers[-2][1])
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
c
8
rgb hashed: 879923da020d1533f4d8e921ea7bac61e8ba41d3c89d17a4d14e3a89c6780d5d
```

<!-- grade_id: cell-a17e088c506ff934, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '879923da020d1533f4d8e921ea7bac61e8ba41d3c89d17a4d14e3a89c6780d5d', "Wrong answer."
```

### 1.5) Tuple of size one

How can we create a tuple of size 1 __without__ using the function __tuple__?

a) (5)   
b) 5   
c) [5]   
d) 5,   

<!-- grade_id: cell-d239485262f1edcd, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
### BEGIN SOLUTION
answer = "d"
print(answer)
this_tuple = 5,
print(type(this_tuple))
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
d
<class 'tuple'>
rgb hashed: 3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633
```

<!-- grade_id: cell-3e575a27a20113ad, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633', "Wrong answer."
```

###  1.6) Replace values in a tuple

Can a tuple be modified after its creation?

<!-- grade_id: cell-1d7ed0f610a63b1a, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "yes"
#answer = "no"
### BEGIN SOLUTION
answer = "no"
print(answer)
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
no
rgb hashed: 04a06452677210a3cdaec376fd5ebbca1714cb7af9e62bf5cce1644310a9086a
```

<!-- grade_id: cell-c7188110ce71150e, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '04a06452677210a3cdaec376fd5ebbca1714cb7af9e62bf5cce1644310a9086a', "Wrong answer."
```

### 1.7) Merge two tuples

Considering the following tuples:

```python
left = (1,11,22215,7,14,1,11,9,1,6,2,5)
right = (1,24,50,45,2,45,1,1,2,1,2,1,88,9,9,9,44,5,2)
```

Create a tuple by merging tuples above and assign it to a third tuple called `this_tuple`. Elements in tuple __`left`__ should come before the ones on the __`right`__ tuple.

<!-- grade_id: cell-05376cff3e8e7950, points: 0, grade: false, solution: true -->
```python
left = (1,11,22215,7,14,1,11,9,1,6,2,5)
right = (1,24,50,45,2,45,1,1,2,1,2,1,88,9,9,9,44,5,2)
#hint: Using operations between tuples.
#this_tuple = ...

### BEGIN SOLUTION
this_tuple = left + right
print(this_tuple)
print(type(this_tuple))
print(len(this_tuple))
print(this_tuple[-3])
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(this_tuple).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
(1, 11, 22215, 7, 14, 1, 11, 9, 1, 6, 2, 5, 1, 24, 50, 45, 2, 45, 1, 1, 2, 1, 2, 1, 88, 9, 9, 9, 44, 5, 2)
<class 'tuple'>
31
44
rgb hashed: 0c1afd35431992aba438b9382b08332f4d4fed7d4380e538c047e42a223a9dd5
```

<!-- grade_id: cell-74db790f2e9e6836, points: 1, grade: true, solution: false -->
```python
assert isinstance(this_tuple, tuple), "The result should be a tuple."
assert len(this_tuple) == 31, "The merging is not right."
assert this_tuple[-3] == 44, "Re-check the order of the tuples."
assert hashlib.sha256(json.dumps(this_tuple).encode()).hexdigest() == '0c1afd35431992aba438b9382b08332f4d4fed7d4380e538c047e42a223a9dd5'
```

## Exercise 2: Lists  <a name="2"></a> 

This exercise covers topics learned regarding lists.

###  2.1) List Creation

Write a list named `this_list` with length five and with the string `"bananas"` on negative index -4.

<!-- grade_id: cell-f8854927f9656f23, points: 0, grade: false, solution: true -->
```python
#bananas = ...
#this_list = ...

### BEGIN SOLUTION
bananas = "bananas"
this_list = ["kiwi", bananas, "melon", "lemon", "kiwi"]
print(this_list)
print(type(this_list))
print(len(this_list))
print(this_list[-4])
### END SOLUTION
```

Output:
```
['kiwi', 'bananas', 'melon', 'lemon', 'kiwi']
<class 'list'>
5
bananas
```

<!-- grade_id: cell-5de18c0078e98c2e, points: 1, grade: true, solution: false -->
```python
assert isinstance(this_list, list), "The result should be a list."
assert len(this_list) == 5, "The length is not quite right."
assert this_list[-4] == "bananas", "There is no \"bananas\" on position -4."
```

###  2.2) Delete and append values in a list

Giving the list `ice_cream`, delete `"chocolate"` and append `"dulce de leche"`.

<!-- grade_id: cell-c22be3e4356c6e22, points: 0, grade: false, solution: true -->
```python
ice_cream = ["lemon", "stracciatella", "pistacchio", "chocolate", "vanilla"]
### BEGIN SOLUTION
print(len(ice_cream))
del ice_cream[-2]
ice_cream.append("dulce de leche")
print("chocolate" in ice_cream)
print(ice_cream[-1])
### END SOLUTION
```

Output:
```
5
False
dulce de leche
```

<!-- grade_id: cell-298523708d65510d, points: 1, grade: true, solution: false -->
```python
assert len(ice_cream) == 5, "The operations were not performed as asked."
assert "chocolate" not in ice_cream, "There is still \"chocolate\" on the ice_cream."
assert ice_cream[-1] == "dulce de leche", "Did you append the exact ingredient?"
```

###  2.3) Delete the last value in a list

Considering the following list, `ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]`, what is the right answer if we want to delete `"vanilla"` in the last position?

a) `del ice_cream[1]`   
b) `del ice_cream[-1]`    
c) `ice_cream.remove("vanilla")`   
d) `ice_cream[-1] = None`   

<!-- grade_id: cell-fafe1f47881998b6, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"

### BEGIN SOLUTION
answer = "b"

ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]
del ice_cream[1]
print(ice_cream)

ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]
del ice_cream[-1] 
print(ice_cream)

ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]
ice_cream.remove("vanilla")
print(ice_cream)

ice_cream = ["lemon", "vanilla", "stracciatella", "pistacchio", "chocolate", "vanilla"]
ice_cream[-1] = None
print(ice_cream)


print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
['lemon', 'stracciatella', 'pistacchio', 'chocolate', 'vanilla']
['lemon', 'vanilla', 'stracciatella', 'pistacchio', 'chocolate']
['lemon', 'stracciatella', 'pistacchio', 'chocolate', 'vanilla']
['lemon', 'vanilla', 'stracciatella', 'pistacchio', 'chocolate', None]
rgb hashed: c100f95c1913f9c72fc1f4ef0847e1e723ffe0bde0b36e5f36c13f81fe8c26ed
```

<!-- grade_id: cell-6589184220a059bc, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == 'c100f95c1913f9c72fc1f4ef0847e1e723ffe0bde0b36e5f36c13f81fe8c26ed', "Wrong answer."
```

###  2.4) List Operations

Using list operations, create a list called `hello_world` of size 10000. It should have just two unique values, `"hello"` and `"world"`. Starting with `"hello"` then `"world"`, then `"hello"` and so on.

<!-- grade_id: cell-0989f9a6a1c1f63e, points: 0, grade: false, solution: true -->
```python
#Hint: multiply a list by a int 
#hello_world = ...

### BEGIN SOLUTION
hello_world = ["hello", "world"] * 5000
print(len(hello_world))
print(hello_world[0])
print(hello_world[3])
print(hello_world[4532])
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(hello_world).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
10000
hello
world
hello
rgb hashed: ba01ef9ff3158651119ad5f13ef61b559dd73eff752f4a2cc05e3912c2c85d32
```

<!-- grade_id: cell-4ac5d3b434c8a490, points: 2, grade: true, solution: false -->
```python
assert len(hello_world) == 10000, "The size is not right."
assert hello_world[0] == "hello", "I guess you started on the wrong foot."
assert hello_world[3] == "world", "Are the words alternating?"
assert hello_world[4532]== "hello", "Are the words alternating?"
```

###  2.5) Replace values and sort a list

Replace `"pistacchio"` in the list `ice_cream` by `"cream"`. Add the sub-list `others` to the end of the list `ice_cream`.   
After these operations, sort the elements in the list and convert it to a __tuple__.

<!-- grade_id: cell-7327a77943c8fe67, points: 0, grade: false, solution: true -->
```python
ice_cream = ["lemon", "stracciatella", "pistacchio", "chocolate", "vanilla"]
others = ["dulce de leche", "caramel", "cookies", "peanut butter"] * 100
### BEGIN SOLUTION
ice_cream[2] = "cream"
ice_cream = ice_cream + others
ice_cream.sort()
ice_cream = tuple(ice_cream)
print(len(ice_cream))
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(ice_cream).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
405
rgb hashed: bbe23485294172ad0fab1f37eee62d4f518a296aad06ea944d6d7a31e6ef58a6
```

<!-- grade_id: cell-955f097b2db82b73, points: 1, grade: true, solution: false -->
```python
assert len(ice_cream) == 405, "The lists weren't merged right."
assert ice_cream.index("cream") == 201, "Did you replaced \"pistacchio\"?"
assert isinstance(ice_cream, tuple), "the result is not a tuple."
assert hashlib.sha256(json.dumps(ice_cream).encode()).hexdigest() == 'bbe23485294172ad0fab1f37eee62d4f518a296aad06ea944d6d7a31e6ef58a6'
```

## Exercise 3: Dictionaries  <a name="3"></a>

This exercise covers topics learned regarding dictionaries.

###  3.1) Create a dictionary

Create a dictionary called `this_dict` with 5 key-value pairs where the keys are strings and values are lists.

<!-- grade_id: cell-bb5b978acdaaae16, points: 0, grade: false, solution: true -->
```python
#this_dict = ...

### BEGIN SOLUTION
this_dict = {"a":[],"b":[],"c":[],"d":[],"e":[]}
print(len(this_dict))
print(isinstance(list(this_dict.keys())[0], str))
print(isinstance(list(this_dict.values())[0], list))
print(isinstance(list(this_dict.keys())[2], str))
### END SOLUTION
```

Output:
```
5
True
True
True
```

<!-- grade_id: cell-0f3a2761c0922b5f, points: 1, grade: true, solution: false -->
```python
assert isinstance(this_dict, dict), "The result is not a dictionary."
assert len(this_dict) == 5, "The dictionary doesn't have 5 key-value pairs."
assert isinstance(list(this_dict.keys())[0], str), "The dictionary keys are not strings."
assert isinstance(list(this_dict.values())[0], list), "The dictionary values are not lists."
```

### 3.2) Extract a value from a dictionary

Considering the following dictionary named `groceries`:

```python
groceries = {             
             "bread": {"type": "grains", "price_per_unit": 2, "quantity_purchased": 1},   
             "onions": {"type": "vegetables", "price_per_unit": 0.5, "quantity_purchased": 2},   
             "spinages": {"type": "vegetables" , "price_per_unit": 1.5, "quantity_purchased": 1}   
            }
```

What is the notation that we should use in order to extract "grains"?

a) `groceries["type"]`   
b) `groceries["bread"]`   
c) `groceries[0][0]`   
d) `groceries["bread"]["type"]`

<!-- grade_id: cell-fd6a36d5f663708b, points: 0, grade: false, solution: true -->
```python
#uncomment the right answer
#answer = "a"
#answer = "b"
#answer = "c"
#answer = "d"
### BEGIN SOLUTION
answer = "d"
print(answer)
groceries = {   
             "bread": {"type": "grains", "price_per_unit": 2, "quantity_purchased": 1},   
             "onions": {"type": "vegetables", "price_per_unit": 0.5, "quantity_purchased": 2},   
             "spinages": {"type": "vegetables" , "price_per_unit": 1.5, "quantity_purchased": 1}   
            }
print(groceries["bread"]["type"])
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(answer).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
d
grains
rgb hashed: 3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633
```

<!-- grade_id: cell-978adad567f2a3cb, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(answer).encode()).hexdigest() == '3fa5834dc920d385ca9b099c9fe55dcca163a6b256a261f8f147291b0e7cf633', "Wrong answer."
```

###  3.3) Replace, Append and Delete operations on dictionaries

Considering the following dictionary named groceries:

```python
groceries = {   
             "bread": {"type": "grains", "price_per_unit": 2, "quantity_purchased": 1},   
             "onions": {"type": "vegetables", "price_per_unit": 0.5, "quantity_purchased": 2},   
             "spinages": {"type": "vegetables" , "price_per_unit": 1.5, "quantity_purchased": 1}   
            }
```

Complete the following questions regarding Replace, Append and Delete of key-value pairs in a dictionary.

####  3.3.1) Replace a value on a key-value pair

Considering the dictionary `groceries`, update the `price_per_unit` for bread from 2 to 3.   

<!-- grade_id: cell-095cb4bce1528b15, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
groceries["bread"]["price_per_unit"] = 3
print(groceries["bread"]["price_per_unit"])
### END SOLUTION
```

Output:
```
3
```

<!-- grade_id: cell-c5c653d824ffcda9, points: 1, grade: true, solution: false -->
```python
assert groceries["bread"]["price_per_unit"]==3
```

####  3.3.2) Add a key-value pair

Add rice to our groceries. It should be of type grains, with a price per unit of 1 and in quantity 2.

<!-- grade_id: cell-bacb5fca35772b35, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
groceries["rice"]={"type": "grains", "price_per_unit": 1, "quantity_purchased": 2}
print(groceries["rice"])
### END SOLUTION
```

Output:
```
{'type': 'grains', 'price_per_unit': 1, 'quantity_purchased': 2}
```

<!-- grade_id: cell-59704708ceb3f98c, points: 1, grade: true, solution: false -->
```python
assert groceries["rice"]["type"] == "grains", "The type of grocery is wrong."
assert groceries["rice"]["price_per_unit"] == 1, "The price per unit is wrong."
assert groceries["rice"]["quantity_purchased"] == 2, "The quantity purchased is wrong."
```

####  3.3.3) Delete a key-value pair

Delete onions from our `groceries` dictionary.

<!-- grade_id: cell-c8c6802714175503, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
del groceries["onions"]
"onions" not in groceries.keys()
### END SOLUTION
```

Output:
```
True
```

<!-- grade_id: cell-9b202febf2268e6d, points: 1, grade: true, solution: false -->
```python
assert "onions" not in groceries.keys(), "There's still \"onions\" on the groceries."
```

###  3.4) Extract keys and values manually

Considering the following `people_age` dictionary, write down a list with all the keys in the dictionary and another one with all the values. Even if a value is duplicated, write it as many times as it appears. Return the lists __sorted__.

`people_age = {"joao":"senior", "bernardo":"adult", "joao":"child", "alberto":"adult", "amilcar":"child" }`

Create the lists manually (without using methods) by just looking at the dictionary above and writing down the values on lists. Also, try to not print the output of the dictionary `people_age`. The idea of this exercise is to guess the output of the dictionary in terms of keys and values.

<!-- grade_id: cell-1e7bf909ee0d2ab8, points: 0, grade: false, solution: true -->
```python
#keys = ...
#values = ...
### BEGIN SOLUTION
keys = ["joao", "bernardo", "alberto", "amilcar"]
values = ["child", "adult", "adult", "child"]
keys.sort()
values.sort()
print(len(keys))
print(len(values))
print(keys)
print(values)
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(keys).encode()
    ).hexdigest()
)
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(values).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
4
4
['alberto', 'amilcar', 'bernardo', 'joao']
['adult', 'adult', 'child', 'child']
rgb hashed: 73a6b69681d6d29dc146cfba3fbeed0cf41d233e0ef7ed0cdda84bb0e7c64924
rgb hashed: 39fb994d286387448ffae1e44256baace4d8329d6bd59875aeee5d6b7d17b291
```

<!-- grade_id: cell-89fa0ee93c6519f4, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(json.dumps(keys).encode()).hexdigest() == '73a6b69681d6d29dc146cfba3fbeed0cf41d233e0ef7ed0cdda84bb0e7c64924', "The keys list is incorrect."
assert hashlib.sha256(json.dumps(values).encode()).hexdigest() == '39fb994d286387448ffae1e44256baace4d8329d6bd59875aeee5d6b7d17b291', "The values list is incorrect."
```

###  3.5) Extract keys and values using methods

Considering the following dictionary called `people_age`, extract keys and values of the dictionary to a variable called `names` and `ages`, respectively. Convert the variable `ages` to a list called `list_ages` and calculate how many people are senior and assign the value to `n_seniors`.   
Contrary to exercise 3.4, on this one you should use methods to extract keys and values.

<!-- grade_id: cell-1f5f20a99e9e053b, points: 0, grade: false, solution: true -->
```python
#Hint: find out what is the type of the variable age, you might need to convert this variable to another type.
#Hint: you might need to use your list knowledge.
people_age = {"joao":"senior", "bernardo":"adult", "gabriel":"child", "antonio":"senior", "maria":"senior", "joel":"adult", "ines":"adult", "alberto":"adult", "amilcar":"senior", "emilia":"adult", "ana":"adult", "margarida":"adult" }
#names = ...
#ages = ...
#list_ages = ..
#n_seniors = ...

### BEGIN SOLUTION
names = people_age.keys()
print(names)
print(len(names))
ages = people_age.values()
print(ages)
print(len(ages))
list_ages = list(ages)
n_seniors = list_ages.count("senior")

print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(list(names)).encode()
    ).hexdigest()
)
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(list(ages)).encode()
    ).hexdigest()
)
print(
    'rgb hashed:',
    hashlib.sha256(
        json.dumps(n_seniors).encode()
    ).hexdigest()
)
### END SOLUTION
```

Output:
```
dict_keys(['joao', 'bernardo', 'gabriel', 'antonio', 'maria', 'joel', 'ines', 'alberto', 'amilcar', 'emilia', 'ana', 'margarida'])
12
dict_values(['senior', 'adult', 'child', 'senior', 'senior', 'adult', 'adult', 'adult', 'senior', 'adult', 'adult', 'adult'])
12
rgb hashed: 9a2a381eef2e8212d8c8b48ab6531a62c064e9fbc39091015c8131138174eaf1
rgb hashed: 7ec678f76d3dec3621b7f64343d89d6d9a0b8faaea6b62c529b937439e3b93f9
rgb hashed: 4b227777d4dd1fc61c6f884f48641d02b4d121d3fd328cb08b5531fcacdabf8a
```

<!-- grade_id: cell-2977c2eb6bba3963, points: 1, grade: true, solution: false -->
```python
assert len(names) == 12, "The number of names is incorrect."
assert hashlib.sha256(json.dumps(list(names)).encode()).hexdigest() == '9a2a381eef2e8212d8c8b48ab6531a62c064e9fbc39091015c8131138174eaf1', "The variable names is incorrect."
assert len(ages) == 12, "The number of ages is incorrect."
assert hashlib.sha256(json.dumps(list(ages)).encode()).hexdigest() == '7ec678f76d3dec3621b7f64343d89d6d9a0b8faaea6b62c529b937439e3b93f9', "The variable ages is incorrect."
assert isinstance(list_ages, list), "list_ages should be a list."
assert hashlib.sha256(json.dumps(n_seniors).encode()).hexdigest() == '4b227777d4dd1fc61c6f884f48641d02b4d121d3fd328cb08b5531fcacdabf8a', "The number of seniors is incorrect."
```